Modulation of aperiodic EEG activity provides sensitive index of cognitive state changes during working memory task

Author response:

The following is the authors’ response to the previous reviews

Reviewer 1:

We thank Reviewer 1 for the discussion on the possible causes of ERPs and their relevance for the interpretation of changes in aperiodic activity. We have changed the relevant paragraph to read as follows: For example, ERPs may reflect changes in periodic activity, such as phase resets (Makeig et al., 2002), or baseline shifts (Nikulin et al., 2007). ERPs may also capture aperiodic activity, either in the form of evoked transients triggered by an event (Shah et al., 2004) or induced changes in the ongoing background signal. This has important implications: evoked transients can alter the broadband spectrum without implying shifts in ongoing background activity, whereas induced aperiodic changes may signal different neural mechanisms, such as shifts in the excitation-inhibition balance (Gao et al., 2017).

Reviewer 1 argued that a time point-by-time point comparison between ERPs and aperiodic parameters may not be the most appropriate approach, since aperiodic time series have lower temporal resolution than ERPs. Reviewer suggested comparing their topographies instead. We had already done this in the first version of the paper (see Fig. S7: https://elifesciences.org/reviewedpreprints/101071v1#s10). However, in the second version, we opted to use linear mixed models for each channel-time point in order to maintain consistency with the other analyses in the paper (e.g. the comparison between FOOOF parameters and baseline-corrected power).

Nevertheless, we repeated the topographic correlations as in the first version, and the results are shown below. Correlations were computed for each time point, subject and condition, and then averaged across these dimensions for visualisation. The pattern differs from that of the linear mixedmodel results (see Fig. S14), with notable correlations appearing after ~0.5 s for the exponent and after ~1.0 s for the offset. Still, the correlations remain low, suggesting that aperiodic parameters and ERPs encode different information (at least in this dataset).

Additionally, to control for the effect of smearing we have performed the same linear mixed model analysis as in Fig. S14 on low-pass filtered ERPs (with cut-off 10 Hz), and the results were largely similar as in Fig. S14.

Reviewer 1 discussed two possible explanations for the observed correlations between baselinecorrected power and FOOOF parameters (Figure 4): “The correlation between the exponent and lowfrequency activity could be of either direction: low frequency power changes could reflect 1/f shifts, or exponent estimates might be biased by undetected delta/theta activity. I think that one other piece of evidence /…/ to intuitively highlight why the latter is more likely is the /…/ decrease at high ("transbeta") frequencies, which suggests a rotational shift /../.” We agree with the interpretation that lowfrequency power changes in our data primarily reflect 1/f shifts. However, we are uncertain about the reviewer’s statement that the “latter” explanation (i.e., bias in exponent estimates due to delta/theta activity) is more likely. Given the context, we believe the reviewer may have intended to say the “former” explanation is more likely.

We agree with the reviewers' observation that rhythmicity, as estimated using the pACF, can be independent of power (Myrov et al., 2024, Fig. 1). However, it seems that in real (non-simulated) datasets, the pACF and power spectral density (PSD) are often moderately correlated (e.g. Myrov et al., 2024, Fig. 5).

Reviewer 1 asked whether we had examined aperiodic changes in the data before and after subtracting the response-locked ERPs. We did not carry out this extra analysis as, as the reviewer suggests, it would have been excessive – the current version of the paper already contains more than 60 figures. As mentioned in the manuscript, we acknowledge the possibility that response-locked ERPs contribute to the second aperiodic component. However, due to the weak correlation between reaction times and aperiodic activity, the presence of both components throughout the entire epoch (in at least the first and third datasets) and the distinct differences between the ERPs and the aperiodic activity in the different conditions (see Fig. 8 vs. Fig. S13), we cannot conclusively determine whether the second aperiodic component is directly related to motor responses. Finally, we agree with the reviewer that the distribution of the response-locked ERP more closely resembles the frontocentral (earlier) aperiodic component than the later post-response component. We have amended the relevant paragraph in the Discussion to include these observations. ”While it is possible that response-related ERPs contributed to the second aperiodic component, several observations suggest otherwise: both aperiodic components were present throughout the entire epoch, differences between conditions diverged between ERPs and aperiodic activity (compare Figure 8 and Figure S16), and the associations with reaction times were weak. Moreover, the distribution of the response-locked ERP qualitatively resembled the earlier frontocentral aperiodic component more than the later post-response component. Taken together, these findings suggest that ERPs and aperiodic activity capture distinct aspects of neural processing, rather than reflecting the same underlying phenomenon.”

We agree with Reviewer 1 that our introduction of aperiodic activity was abrupt, and that the term 'aperiodic exponent' required definition. We have now defined it as the spectral steepness in log–log space (i.e. the slope), and have added a brief explanatory sentence to the introduction.

Reviewer 1 noted that the phrase 'task-related changes in overall power' could be misinterpreted as referring to total (broadband) power, and recommended that we specify a frequency range. We agree, so we have replaced 'overall power' with 'spectral power within a defined frequency range'.

We agree with Reviewer 1 that the way we worded things in the Discussion section regarding alpha activity and inhibitory processes was awkward and could easily be misread. We have rephrased the sentences and added a brief explanation to avoid implying a direct link between alpha attenuation and neural inhibition.

Furthermore, based on the reviewer’s suggestion, we added a brief comment in the Discussion section (Theoretical and methodological implications) on theoretical perspectives regarding the interaction between age and aperiodic activity.

Reviewer 1 suggested including condition as a fixed effect in order to examine whether the relationship between FOOOF parameters and baseline-corrected power is modulated by condition. Specifically, the reviewer proposed changing our model from

baseline_corrected_power ~ 1 + fooof_parameter + (1|modality) + (1|nback) + (1|stimulus) + (1|subject)

to

baseline_corrected_power ~ 1 + fooof_parameter + modality*nback *stimulus + (1|subject)

While we appreciate this suggestion, we believe that including design variables as fixed effects would confound the interpretation of (marginal) R² as a measure of the association between FOOOF parameters and baseline-corrected power. Our primary question in this analysis was about the fundamental relationship between these measures, not how experimental conditions moderate this relationship.

To address the reviewer's concern regarding condition-specific effects, we conducted separate analyses for each condition using a simpler model:

baseline_corrected_power ~ 1 + fooof_parameter + (1|subject)

The results (now included in the Supplement, Fig. S4–S6) show generally smaller effect sizes compared to our original random-effects model, with notable differences between conditions. The 2-back conditions, particularly the non-target trials, exhibited the weakest associations. Despite these differences, the overall patterns remained consistent with our original findings: exponent and offset exhibited positive associations at low frequencies (delta, theta) and negative associations at higher frequencies (beta, low gamma), while periodic activity correlated substantially with baselinecorrected power in the alpha, beta, and gamma ranges.

However, this condition-specific approach has important limitations. With only 47 subjects per condition, the statistical power is insufficient for stable correlation estimates (Schönbrodt & Perugini, 2013; https://doi.org/10.1016/j.jrp.2013.05.009). This likely explains why the effects are smaller and less stable effects than in our original model, which uses the full dataset's power while appropriately accounting for condition-related variance through random effects. Since these additional analyses do not alter our primary conclusions, we have included them in the Supplement for completeness and made a minor change in the Discussion section.

Reviewer 1 asked what channels are lines on Figure 9 based on. As stated in the Methods section, “We fitted models in a mass univariate manner, that is for each channel, frequency (where applicable), and time point separately. /…/ For the purposes of visualisation, p-values were averaged across channels (for heatmaps or lines) or across time (for topographies).” Therefore, the lines and heatmaps apply to all channels.

Reviewer 2:

We would like to thank reviewer 2 for their detailed explanation of the expected behaviour of the specparam algorithm. We have added the following explanation to the Methods section:

Importantly, as noted by the reviewer, this behaviour reflects an explicit design choice of the algorithm: to avoid overfitting ambiguous peaks at the edges of the spectrum, FOOOF excludes peaks that are too close to the boundaries. This exclusion is controlled by the _bw_std_edge parameter, which defines the distance that a peak must be from the edge in order to be retained (in units of standard deviation; set to 1.0 by default). Therefore, although the algorithm is functioning as intended, users should be careful when interpreting aperiodic parameters in datasets where lowfrequency oscillatory activity might be expected.

In line with the reviewer’s suggestion we have added a version of specparam to the paper.

We thank reviewer 2 for pointing out two studies that used a time-resolved approach to spectral parameterisation. We have updated the text accordingly:

Although a similar approach has been used to track temporal dynamics in sleep and resting state (e.g., Wilson et al., 2022; Ameen et al., 2024), as well as in task-based contexts (e.g., Barrie et al., 1996; Preston et al., 2025), its specific application to working memory paradigms remains underexplored.

Reviewer 3:

Reviewer 3 notes that the revised manuscript feels less intriguing than the original version. While we understand this concern, we believe this difference arises from a misalignment in expectations regarding the scope and purpose of our study. We think the reviewer is interpreting our work as focusing on whether theta activity is elicited in a paradigm that reliably produces theta oscillations. In contrast, our study is framed around a working memory task in which, based on prior literature, we expected to observe theta activity but instead found an absence of theta spectral peaks in almost all participants. Note that the absence of theta is already noteworthy in itself, given that theta oscillations are believed to play a crucial role in working memory.

Importantly, Van Engen et al. (2024) have recently reported similar findings:

”While we did not observe load-dependent aperiodic changes over the frontal midline, we did reveal the possibility that previous frontal midline theta results that do not correct for aperiodic activity likely do not reflect theta oscillations. /…/ While our results do not invalidate previous research into extracranial theta oscillations in relation to WM, they challenge popular and widely held beliefs regarding the mechanistic role for theta oscillations to group or segregate channels of information”.

From this perspective, we maintain that the following statements are still justified:

“substantial portion of the changes often attributed to theta oscillations in working memory tasks may be influenced by shifts in the spectral slope of aperiodic activity”

"Note that although no prominent oscillatory peak in the theta range was observed at the group level, and some of this activity could potentially fall within the delta range, similar lowfrequency patterns have often been referred to as 'theta' in previous work, even in the absence of a clear spectral peak"

These formulations are intended to emphasize existing interpretations of changes in low-frequency power as theta oscillations in related research.

Next, Reviewer 3 pointed out that “spectral reflection (peak?) in spectral power plot does not imply that an event is repeating (i..e. oscillatory).” We agree with the reviewer that not every spectral peak implies a true oscillation. To address this, we complemented the power analyses with a measure of rhythmicity (phase autocorrelation function, pACF) after the first round of reviews, and the pACF results were largely similar to those for periodic activity. These results suggest that, in our case, periodic activity is indeed largely oscillatory.

However, we do agree with the reviewer that the term “oscillatory” is not interchangeable with “periodic”. To address this, we reviewed the paper for all appearances of “oscillations”, “oscillatory” and related terms, and replaced them with “power”, “spectral” or “periodic activity” where appropriate (all changes are marked in red in the latest version of the manuscript).

Examples of corrections:

Changes in aperiodic activity appear as low-frequency oscillations in baseline-corrected time-frequency plots à low-frequency power

“The periodic component includes only the parameterised oscillatory peak” à spectral peak

“FOOOF decomposition may miss low-frequency oscillations near the edges of the spectrum” à low-frequency peaks

We disagree with the reviewer’s assertion that the subtitle “Aperiodic parameters are largely independent of oscillatory activity” is misleading for a methods oriented paper. Namely, the full subtitle is “Rhythmicity analysis reveals aperiodic parameters are largely independent of oscillatory activity”. Since rhythmicity is a phase-based measure that requires repeating dynamics and is therefore indicative of oscillations, we believe this phrasing is technically accurate.

Finally, we would like to emphasise our contribution once again. Our analyses of rhythmicity, spectrally parameterised power, and baseline-corrected power offer different perspectives on the data. Each of these analyses may lead to different interpretations, but performing all of them on the same data provides a more comprehensive insight into what is actually going on in the data.

Our findings demonstrate that conclusions drawn from a single analytical approach may be incomplete or misleading. For example, as we discuss in the paper, many studies examine thetagamma coupling in scalp EEG during n-back tasks without first establishing whether theta activity genuinely oscillates (e.g. Rajji et al., 2016). The absence of true theta oscillations would undermine the validity of such analyses. Our multifaceted approach provides researchers with a systematic framework for validating oscillatory assumptions before proceeding with more complex analyses.

Changes from the earlier version to the revised version

Abstract

To better understand cognitive processes, it is valuable to consider both periodic (oscillatory) and aperiodic electrophysiological brain activity. In this study, we aim to clarify how the periodic and aperiodic electrophysiological components respectively reflect the cognitive processes involved in working memory. Fifty-seven participants performed an n-back task while their brain activity was recorded using EEG. The analysis of both components of the EEG spectrum during task performance indicates that both periodic and aperiodic activities exhibit distinct task-related spatiotemporal dynamics that are closely related to cognitive demands. The results suggest that a substantial portion of the changes often attributed to theta oscillations in working memory tasks may be influenced by shifts in the spectral slope of aperiodic activity. This finding indicates that the modulation of aperiodic activity, associated with cognitive control processes, could provide a more sensitive index of cognitive state changes than previously recognised. To further confirm our findings, we also used these analysis methods in an item-recognition task, which showed similar patterns of periodic and aperiodic activity. These observations challenge the conventional understanding of low-frequency oscillations in cognitive processing and raise concerns about the routine practice of EEG baseline correction in time-frequency analysis, as it may obscure significant modulations in continuous working memory tasks. Consequently, the inclusion of aperiodic activity as a fundamental component of EEG analysis is likely to be critical for the accurate representation of the neural basis of cognition.

Introduction

Working memory is a crucial cognitive ability that enables goal-directed behaviour and higher cognitive function. It involves the maintenance and manipulation of information in the absence of sensory stimuli (Baddeley, 2012) and depends on a complex set of cognitive processes and neural mechanisms (Wager and Smith, 2003). Due to its central role in cognition, working memory has been extensively studied using electroencephalography (EEG). This method allows for time-frequency analysis of brain activity (Cohen, 2014), often focusing on predefined frequency bands such as theta (4–8 Hz), alpha (6–12 Hz), and beta (13–30 Hz). One widely used paradigm for investigating working memory via time-frequency analysis is the n-back task (Owen et al., 2005; Gevins et al., 1997), which integrates key working memory processes such as encoding, maintenance, updating, and recall (Pesonen et al., 2007; Chen and Huang, 2016).

Earlier version: Beyond band-specific oscillations, EEG power spectra exhibit a broadband decrease in power with increasing frequency, following a 1/f scaling pattern (where f is frequency). This pattern is widely thought to reflect non-oscillatory (aperiodic) neural activity that lacks distinct spectral peaks. Because aperiodic activity spans a continuous frequency range it can obscure changes in periodic (oscillatory) activity, complicating interpretation, unless appropriately separated or accounted for (Cohen, 2014; Donoghue et al., 2020b).

Revised version: Beyond band-specific spectral modulations, EEG power spectra exhibit a broadband decrease in power with increasing frequency, following a 1/f scaling pattern (where f is frequency). This pattern is widely thought to reflect non-oscillatory (aperiodic) neural activity that lacks distinct spectral peaks. Because aperiodic activity spans a continuous frequency range it can obscure changes in periodic (oscillatory) activity, complicating interpretation, unless appropriately separated or accounted for (Cohen, 2014; Donoghue et al., 2020b).

Earlier version: Given that aperiodic activity is sometimes regarded as background noise, a common approach in task-based EEG studies is to apply baseline correction to emphasize task-related oscillatory changes. Baseline correction aims to normalise the data by removing this background activity, often using decibel or z-score normalisation. These methods typically assume stable aperiodic activity across time (Gyurkovics et al., 2021; Merkin et al., 2023; Donoghue et al., 2020b), however, in continuous paradigms like the n-back task, where no neutral pre-stimulus interval exists, this assumption may not hold. Moreover, recent studies show that ignoring task-related aperiodic fluctuations can distort oscillatory measurements and mask meaningful physiological and behavioural information (Gyurkovics et al., 2021; Donoghue et al., 2020a; Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022; Kałamała et al., 2024; Akbarian et al., 2023b).

Revised version: Given that aperiodic activity is sometimes regarded as background noise, a common approach in task-based EEG studies is to apply baseline correction to emphasize task-related changes in power. Baseline correction aims to normalise the data by removing this background activity, often using decibel or z-score normalisation. These methods typically assume stable aperiodic activity across time (Gyurkovics et al., 2021; Merkin et al., 2023; Donoghue et al., 2020b), however, in continuous paradigms like the n-back task, where no neutral pre-stimulus interval exists, this assumption may not hold. Moreover, recent studies show that ignoring task-related aperiodic fluctuations can distort estimates of periodic activity and mask meaningful physiological and behavioural information (Gyurkovics et al., 2021; Donoghue et al., 2020a; Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022; Kałamała et al., 2024; Akbarian et al., 2023b).

Earlier version: Computational models suggest that these aperiodic fluctuations reflect the balance between excitation and inhibition (E:I) in the brain (Gao et al., 2017; Lim and Goldman, 2013), essential for efficient cognitive processing (Ahmad et al., 2022). Aperiodic activity is often approximated as linear in log-log space (Donoghue et al., 2020b; Barry and De Blasio, 2021), and while sometimes referred to as the ‘aperiodic exponent,’ we use the term ‘slope’ here, except where methodological descriptions require exponent notation. A steeper aperiodic slope implies higher low-frequency power and lower high-frequency power, potentially indicating stronger cognitive engagement (Gao et al., 2017; Akbarian et al., 2023a), whereas a flatter slope can reflect more uniform power across frequencies (Gyurkovics et al., 2022; Kałamała et al., 2024). However, some tasks show flatter slopes even under increased engagement (Waschke et al., 2021; Cunningham et al., 2023). Moreover, task-related changes in overall power, typically interpreted as oscillatory dynamics, may be significantly co-modulated by the aperiodic slope shift (Akbarian et al., 2023b; Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022).

Revised version: Computational models suggest that these aperiodic fluctuations reflect the balance between excitation and inhibition (E:I) in the brain (Gao et al., 2017; Lim and Goldman, 2013), essential for efficient cognitive processing (Ahmad et al., 2022). Aperiodic activity is often approximated as linear in log-log space (Donoghue et al., 2020b; Barry and De Blasio, 2021). In the commonly used 1/f formulation, the aperiodic exponent describes how quickly power decreases as frequency increases. When both axes are shown on a logarithmic scale, the line’s slope is the negative of this exponent, so a larger exponent yields a more steeply descending slope. While sometimes referred to as the ‘aperiodic exponent,’ we use the term ‘slope’ here, except where methodological descriptions require exponent notation. A steeper aperiodic slope (larger exponent; more negative slope) implies higher low-frequency power and lower high-frequency power, potentially indicating stronger cognitive engagement (Gao et al., 2017; Akbarian et al., 2023a), whereas a flatter slope (less negative exponent) can reflect more uniform power across frequencies (Gyurkovics et al., 2022; Kałamała et al., 2024). However, some tasks show flatter slopes even under increased engagement (Waschke et al., 2021; Cunningham et al., 2023). Moreover, task-related changes in spectral power within a defined frequency range, typically interpreted as oscillatory dynamics, may be significantly co-modulated by the aperiodic slope shift (Akbarian et al., 2023b; Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022).

Earlier version: Given these findings, it is methodologically imperative to carefully parameterise both periodic and aperiodic components of an EEG before interpreting time-frequency analysis results. Common approaches to spectral or time-frequency analysis often assume that distinct peaks in the power spectrum reflect true oscillatory activity once the 1/f background is removed (Donoghue et al., 2020b,a). However, the presence of power in a particular frequency band is not necessarily indicative of oscillatory dynamics, as aperiodic activity can constitute a significant proportion of observed power spectra (Bullock et al., 2003; Donoghue et al., 2020a). Moreover, time-frequency representations (TFRs) typically highlight temporal fluctuations in power (Cohen, 2014), but do not inherently differentiate between periodic and aperiodic contributions. If part of the TFR signal arises from aperiodic shifts, baseline correction may confound rather than clarify the nature of observed effects.

Revised version: Given these findings, it is methodologically imperative to carefully parameterise both periodic and aperiodic components of an EEG before interpreting time-frequency analysis results. Common approaches to spectral or time-frequency analysis often assume that distinct peaks in the power spectrum reflect true oscillatory activity once the 1/f background is removed (Donoghue et al., 2020b,a). However, the presence of power in a particular frequency band is not necessarily indicative of oscillatory dynamics, as aperiodic activity can constitute a significant proportion of observed power spectra (Bullock et al., 2003; Donoghue et al., 2020a, 2022). Moreover, time-frequency representations (TFRs) typically highlight temporal fluctuations in power (Cohen, 2014), but do not inherently differentiate between periodic and aperiodic contributions. If part of the TFR signal arises from aperiodic shifts, baseline correction may confound rather than clarify the nature of observed effects.

In this study, we examined how aperiodic and periodic EEG components each contribute to working memory demands in the n-back task, addressing potential biases introduced by baseline correction. Furthermore, we focused on disentangling how aperiodic slope is influenced by memory load, stimulus type, or task modality. We based our analyses on a dataset of middle-aged to older adults, half of whom reported subjective cognitive complaints; nonetheless, behavioural outcomes and electrophysiological measures did not differ between groups. To assess the generalisability of our findings, we applied the same analysis to two additional datasets which allowed us to test whether similar aperiodic–periodic dynamics emerge across different working memory contexts.

Earlier version: Our study builds on previous research by using high temporal resolution data to investigate the aperiodic slope across two modalities (visuospatial/verbal). In contrast to n-back studies that use discrete time windows (Akbarian et al., 2023b; Kałamała et al., 2024), we analyse the activity at each time point, capturing shortlived slope fluctuations that could reveal crucial insights into continuous working memory processes. While a similar approach has been used to track temporal dynamics in sleep and resting state (e.g., Wilson et al., 2022; Ameen et al., 2024), we are not aware of any study using time-resolved spectral parameterisation in task-based paradigms. While some evidence suggests modality-specific aperiodic variations (Waschke et al., 2021), we propose that the n-back task primarily recruits domain-general processes (Owen et al., 2005; Gazzaley and Nobre, 2012; Olesen et al., 2004), leading to similar slope modulations across both verbal and visuospatial modalities (Gyurkovics et al., 2022). We therefore predict that the aperiodic slope would be modulated by the processing demands of the n-back task, varying with load × and stimulus type but not between visuospatial and verbal conditions.

Revised version: Our study builds on previous research by using high temporal resolution data to investigate the aperiodic slope across two modalities (visuospatial/verbal). In contrast to studies that use discrete time windows (Akbarian et al., 2023b; Kałamała et al., 2024), we analyse the activity at each time point, capturing short-lived slope fluctuations that could reveal crucial insights into continuous working memory processes. Although a similar approach has been used to track temporal dynamics in sleep and resting state (e.g., Wilson et al., 2022; Ameen et al., 2024), as well as in task-based contexts (e.g., Barrie et al., 1996; Preston et al., 2025), its specific application to working memory paradigms remains underexplored. While some evidence suggests modality-specific aperiodic variations (Waschke et al., 2021), we propose that the n-back task primarily recruits domain-general processes (Owen et al., 2005; Gazzaley and Nobre, 2012; Olesen et al., 2004), leading to similar slope modulations across both verbal and visuospatial modalities (Gyurkovics et al., 2022). We therefore predict that the aperiodic slope would be modulated by the processing demands of the n-back task, varying with load and stimulus type but not between visuospatial and verbal conditions.

By explicitly separating aperiodic and periodic components and tracking their dynamics over time, our study offers a novel perspective on the relationship between baseline correction and aperiodic activity in continuous tasks.

Results

We analysed the EEG data from the visuospatial and verbal n-back tasks. We transformed the data into the time-frequency domain by convolving it with superlets (Moca et al., 2021). Subsequently, we applied the FOOOF (Fitting Oscillations and One-Over-F) algorithm (Donoghue et al., 2020b) to decompose the activity into periodic and aperiodic components (Figure 2). We estimated two aperiodic parameters using FOOOF: the offset and the exponent. While the FOOOF framework refers to this parameter as the exponent, we use the term ‘aperiodic slope’ throughout the manuscript for consistency with broader literature on spectral power scaling. Note that the exponent is numerically positive, whereas the slope is its negative counterpart (i. e., a steeper (more negative) slope in log-log space corresponds to a higher (positive) exponent value).

Removed in the revised version: Schematic representation of the n-back task.

Removed in the revised version: The n-back working memory task was conducted in two distinct modalities: visuospatial and verbal. In the 2-back condition, participants had to identify whether the current stimulus matched the one presented two steps previously. In the visuospatial modality, the target was a spatial location, whereas in the verbal modality, the target was a letter. In the 0-back condition participants’ task was to respond to a predefined target, with the type of target corresponding to the task modality. Target stimuli are highlighted in orange.

Removed in the revised version: Schematic of the analysis.

Removed in the revised version: The time domain data were first transformed into the time-frequency domain by convolution with superlets (Moca et al., 2021). Next, the periodic and aperiodic components of the power spectrum density were estimated for each time point using FOOOF (Fitting Oscillations and One-Over-F) algorithm (Donoghue et al., 2020b). The aperiodic component was characterised by the aperiodic slope (the negative counterpart of the exponent parameter) and the offset, which together describe the underlying broadband spectral shape.

Earlier version: Changes in aperiodic activity appear as low-frequency oscillations in baseline-corrected time-frequency plots

Revised version: Changes in aperiodic activity appear as low-frequency power in baseline-corrected time-frequency plots

Earlier version: As the first step of the analysis, we compared periodic and aperiodic activity with baseline-corrected time-frequency activity. In the baseline-corrected time-frequency plot, we observed a decrease in alpha and beta power up to 1 second post-stimulus, while low-frequency power increased from 0 seconds post-stimulus (Figure 3, Figure S18). The observed pattern was minimally affected by the choice of baseline (Figure S2). Specifically, whether we employed subtractive or divisive baseline normalisation (Gyurkovics et al., 2021), the qualitative pattern of alpha, beta, and low-frequency (delta and theta) changes remained consistent.

Revised version: As the first step of the analysis, we compared periodic and aperiodic activity with baseline-corrected time-frequency activity. In the baseline-corrected time-frequency plot, we observed a decrease in alpha and beta power up to 1 second post-stimulus, while low-frequency power increased from 0 seconds post-stimulus (Figure 3, Figure S21). The observed pattern was minimally affected by the choice of baseline (Figure S2). Specifically, whether we employed subtractive or divisive baseline normalisation (Gyurkovics et al., 2021), the qualitative pattern of alpha, beta, and low-frequency (delta and theta) changes remained consistent.

Removed in the revised version: Effect of baseline correction and FOOOF decomposition on time-frequency decomposition in the n-back task.

Removed in the revised version: Illustration of the effect of baseline correction and FOOOF decomposition on EEG time-frequency analysis. (A) “Raw” total power from time-frequency decomposition is difficult to interpret due to the 1/f power scaling of EEG power spectra (here representing average data across subjects, channels and conditions). In panel B we applied baseline correction (decibel conversion) using a pre-stimulus interval of -0.5 to -0.2 seconds for comparison. The baseline correction showed a significant decrease in alpha and beta power from 0 to 1 second and a concomitant increase in low-frequency power lasting up to 2 seconds. The observed changes were very similar across different choices of baseline correction (Figure S2). However, it’s unclear from the baseline-corrected data whether the observed changes in the low-frequency range reflect oscillatory or aperiodic contributions. To disentangle these components, we decomposed the time-frequency signal into periodic and aperiodic contributions using spectral parameterisation. (C) The periodic component includes only the parameterised oscillatory peaks, reconstructed from Gaussian fits to the power spectrum (see Methods for details). (D) The aperiodic component reflects the 1/f-like background activity. This decomposition suggests that changes in power in the low-frequency (delta and theta) range may largely reflect changes in aperiodic activity. See Figure S1 for the corresponding figure with a logarithmic y-axis.

Earlier version: Note that although no prominent oscillatory peak in the theta range was observed at the group level, and some of this activity could potentially fall within the delta range, similar low-frequency patterns have often been referred to as ‘theta’ in previous work, even in the absence of a clear spectral peak (e.g., Palomäki et al., 2012; Pesonen et al., 2007; Rossi et al., 2023). To acknowledge both perspectives, we use the more neutral term ‘low-frequency activity’ throughout the paper, while noting its potential overlap with activity commonly described as frontal theta.

Revised version: Note that although no prominent spectral peak in the theta range was observed at the group level, and some of this activity could potentially fall within the delta range, similar low-frequency patterns have often been referred to as ‘theta’ in previous work, even in the absence of a clear spectral peak (e.g., Palomäki et al., 2012; Pesonen et al., 2007; Rossi et al., 2023). To acknowledge both perspectives, we use the more neutral term ‘low-frequency activity’ throughout the paper, while noting its potential overlap with activity commonly described as frontal theta.

Earlier version: Notably, when we examined periodic activity, we only observed changes in alpha and beta frequency bands, while increases in the low-frequency range were absent (Figure 6). This suggests that low-frequency changes may, at least in part, reflect aperiodic activity rather than oscillatory processes. To quantitatively assess this observation, we directly compared the FOOOF and baseline-corrected time-frequency decompositions by performing correlations for each channel-frequency-time point (Figure 4, Figure S3). We observed strong positive correlations between periodic power and baseline-corrected power in the alpha and beta ranges across all channels, but only weak correlations in the low-frequency range. In contrast, both the aperiodic exponent and offset parameters showed strong positive correlations with baseline-corrected power in the theta range. In addition, the exponent showed strong negative correlations with activity above 35 Hz, suggesting that higher frequency power decreases as the aperiodic slope steepens.

Revised version: Notably, when we examined periodic activity, we only observed changes in alpha and beta frequency bands, while increases in the low-frequency range were absent (Figure 6). This suggests that low-frequency changes may, at least in part, reflect aperiodic activity rather than periodic (narrow-band) processes. To quantitatively assess this observation, we directly compared the FOOOF and baseline-corrected time-frequency decompositions by performing correlations for each channel-frequency-time point (Figure 4, Figure S3). We observed strong positive correlations between periodic power and baseline-corrected power in the alpha and beta ranges across all channels, but only weak correlations in the low-frequency range. In contrast, both the aperiodic exponent and offset parameters showed strong positive correlations with baseline-corrected power in the theta range. In addition, the exponent showed strong negative correlations with activity above 35 Hz, suggesting that higher frequency power decreases as the aperiodic slope steepens.

Removed in the revised version: Correlations between baseline-corrected time-frequency and FOOOF-decomposed EEG activity

Removed in the revised version: For each FOOOF parameter, we estimated its similarity to baseline-corrected time-frequency EEG activity at each channel-frequency-timepoint using a linear mixed model. We observed strong correlations between periodic activity and baseline-corrected time-frequency activity in the alpha and beta ranges. In contrast, the exponent and offset showed strong correlations in the lowfrequency range, with the exponent also showing negative correlations in the gamma range. These results suggest that changes in aperiodic activity appear as low-frequency oscillations in baseline-corrected time-frequency plots. Note that while R2 is strictly non-negative, we assigned its sign based on the beta co-efficient from the fixed effect in the model to facilitate interpretation. The values presented are averaged across all channels; see Figure S3 for topographical distributions.

Rhythmicity analysis reveals aperiodic parameters are largely independent of oscillatory activity

Earlier version: A possible concern is that the FOOOF decomposition might not be sufficiently sensitive to peaks at the edges of the frequency spectrum, potentially affecting the interpretation of the results. To address this issue, we conducted simulations demonstrating that when periodic components include peaks at low frequencies (around 3 Hz), they may go undetected, thereby inflating estimates of aperiodic activity (see the Supplement for details).

Revised version: A possible concern is that the FOOOF decomposition might not be sufficiently sensitive to peaks at the edges of the frequency spectrum, potentially affecting the interpretation of the results (We thank the reviewers for alerting us to this concern and suggesting an approach to resolve it.). To address this issue, we conducted simulations demonstrating that when periodic components include peaks at low frequencies (around 3 Hz), they may go undetected, thereby inflating estimates of aperiodic activity (see the Supplement for details).

Earlier version: To overcome this limitation, we employed an alternative metric to assess oscillatory activity: the recently developed measure of rhythmicity, the phase autocorrelation function (pACF) (Myrov et al., 2024). The pACF quantifies the temporal stability of phase dynamics and captures oscillatory activity independent of amplitude fluctuations. Consistent results between FOOOF and pACF would strengthen the conclusion that low-frequency oscillations were truly absent from the recordings, rather than simply undetectable by FOOOF. As expected, the observed rhythmicity patterns were largely consistent with periodic power, particularly in the alpha frequency range (Figure 5, Figure S4, Figure S5). In contrast, correlations between aperiodic parameters and pACF were weak in comparison, supporting the notion that the aperiodic parameter estimates remain largely independent of oscillatory processes.

Revised version: To overcome this limitation, we employed an alternative metric to assess oscillatory activity: the recently developed measure of rhythmicity, the phase autocorrelation function (pACF) (Myrov et al., 2024). The pACF quantifies the temporal stability of phase dynamics and captures oscillatory activity independent of amplitude fluctuations. Consistent results between FOOOF and pACF would strengthen the conclusion that low-frequency oscillations were truly absent from the recordings, rather than simply undetectable by FOOOF.

Added in the revised version: As expected, the observed rhythmicity patterns were largely consistent with periodic power, particularly in the alpha frequency range (Figure 5, Figure S7, Figure S8). In contrast, correlations between aperiodic parameters and pACF were weak in comparison, supporting the notion that the aperiodic parameter estimates remain largely independent of oscillatory processes.

Removed in the revised version: Phase autocorrelation function and its correlations with FOOOF parameters.

Removed in the revised version: Phase autocorrelation function (pACF) (Myrov et al., 2024) is an alternative measure of oscillations that is sensitive to rhythmicity rather than the amplitude. Patterns of pACF (A) were largely consistent with periodic activity (B), particularly in the alpha range, supporting the notion that pACF and periodic activity reflect the same underlying processes. In contrast, the correlations with exponent and offset were low (only results for exponent are shown as results for exponent and offset were essentially the same). See Figure S4 and Figure S5 for detailed figures.

Periodic activity changes with working memory load, stimulus type and modality

Earlier version: Next, we examined the differences in periodic activity as a function of task conditions. We observed a decrease in the alpha and beta bands of periodic activity in the 2-back compared to the 0-back condition over the entire trial period and over the entire scalp (Figure 6, Figure 7). Furthermore, we found a greater decrease in alpha and beta power in response to targets compared to non-targets. In addition, we observed a correlation between beta power and reaction times in central channels. This correlation was positive in the early period (up to 0.5 seconds post-stimulus) and negative in the late period (1–2 seconds post-stimulus). Finally, we observed an interaction between modality and stimulus type in the time period from 0.5 to 1 second post-stimulus, reflecting a greater difference between target and non-target stimuli in the verbal task in the alpha and beta bands compared to the visuospatial task (Figure S7).

Revised version: Next, we examined the differences in periodic activity as a function of task conditions. We observed a decrease in the alpha and beta bands of periodic activity in the 2-back compared to the 0-back condition over the entire trial period and over the entire scalp (Figure 6, Figure 7). Furthermore, we found a greater decrease in alpha and beta power in response to targets compared to non-targets. In addition, we observed a correlation between beta power and reaction times in central channels. This correlation was positive in the early period (up to 0.5 seconds post-stimulus) and negative in the late period (1–2 seconds post-stimulus). Finally, we observed an interaction between modality and stimulus type in the time period from 0.5 to 1 second post-stimulus, reflecting a greater difference between target and non-target stimuli in the verbal task in the alpha and beta bands compared to the visuospatial task (Figure S10).

Removed in the revised version: Changes in periodic (oscillatory) activity as a function of time.

Removed in the revised version: (A) The inspection of periodic activity revealed strong activity in the alpha and beta frequency bands, with a sharp decrease at around 0.5 seconds post-stimulus. The power was stronger in the 0-back condition, compared to the 2-back condition (see also Figure 7). (B) Early beta activity was most prominent at occipital and frontal channels. Note that colour scale ranges in panel B differ between frequency ranges.

Removed in the revised version: Results of the linear mixed model analysis of periodic activity for comparison between conditions.

Removed in the revised version: Significant values are highlighted on the heatmap and marked with yellow circles on the topographies. Only factors with significant differences are shown. A p-value was interpreted as significant if it was significant in at least 3 channels or 3 time points (see Methods for details). Note that there are small discrepancies between significant values on heatmaps and topographies due to averaging across time or channels. A significant difference was observed between the 2-back and 0-back conditions, with a reduction in activity in the 2-back condition, particularly in the alpha and beta frequency bands, starting at 0.5 seconds post-stimulus. The differences between stimulus types were most evident from 0.3 to 1 second post-stimulus, with decreased activity for targets compared to non-targets across the whole scalp. Additionally, a smaller effect of the modality × stimulus type interaction was observed from 0.6 to 1 second post-stimulus (see also Figure S7 for detailed visualisation of the interaction). Associations with reaction times were significant in the beta band across central channels. These associations exhibited a positive correlation in the early phase (0 to 0.5 seconds) and a negative correlation in the later phase (1 to 1.5 seconds).

Aperiodic activity has frontal and parietal components

Earlier version: We observed two distinct components (peaks) in the time course of the exponent (Figure 8). The first component peaked at around 0.3 s after stimulus onset and was most pronounced in frontal channels, whereas the second component peaked at around 0.7 s post-stimulus and was most pronounced in parietal channels. The exponent exhibited significant differences between stimulus types (Figure 9). Specifically, the exponent was larger for targets at approximately 0.3 seconds and for non-targets at approximately 0.7 seconds post-stimulus. We observed no other differences between conditions. We found associations with reaction times at approximately 1.4 seconds post-stimulus in the parieto-occipital channels.

Revised version: We observed two distinct components (peaks) in the time course of the exponent (Figure 8). The first component peaked at around 0.3 s after stimulus onset and was most pronounced in frontal channels, whereas the second component peaked at around 0.7 s poststimulus and was most pronounced in parietal channels. The exponent exhibited significant differences between stimulus types (Figure 9). Specifically, the exponent was larger for targets at approximately 0.3 seconds and for non-targets at approximately 0.7 seconds post-stimulus. We observed no other differences between conditions. We found associations with reaction times at approximately 1.4 seconds post-stimulus in the parieto-occipital channels.

Removed in the revised version: Changes in aperiodic activity (exponent, interpreted as aperiodic slope) as a function of time.

Removed in the revised version: (A) We averaged the time course of the aperiodic activity (exponent) over all channels and observed two components (peaks) of the aperiodic activity. The first component peaked around 0.3 s post-stimulus in the frontal channels. The second component peaked at 0.7 s, was stronger in parietal channels and differed between non-target and target conditions (see Figure 9). The time course of the offset parameter was comparable, although the separation between the frontal and parietal components was more pronounced (see Figure S8). The vertical lines in panel A represent mean reaction time modes. The shaded areas represent 95% Cousineau-Morey within-subjects confidence intervals, adjusted so that non-overlapping intervals correspond to statistically significant differences (Cousineau, 2017).

Removed in the revised version: Results of the linear mixed model analysis of aperiodic exponent (interpreted as slope) for comparison between conditions.

Removed in the revised version: Significant values are shown in blue on line plots and marked with yellow circles on topographies. Only factors with significant differences are shown. The only significant differences between conditions were observed between target and non-target conditions, where the exponent was higher for targets in the early phase (0 to 0.5 s post-stimulus) and for non-targets in the middle phase (0.5 to 1 s post-stimulus). There was also a small association with reaction times in occipital channels around 1.5 seconds post-stimulus. Results were similar for the offset parameter, with an additional effect of load and n-back × stimulus type interaction (Figure S11).

Earlier version: Similar to the exponent, offset also showed a pattern of activity with two components (see Figure S8). In addition, offset was sensitive to differences in load (lower offset for 2-back), and there was also an interaction between load and stimulus type (see Figure S11).

Revised version: Similar to the exponent, offset also showed a pattern of activity with two components (see Figure S11). In addition, offset was sensitive to differences in load (lower offset for 2-back), and there was also an interaction between load × and stimulus type (see Figure S14).

Aperiodic activity does not reflect event-related potentials

Earlier version: To ensure that the aperiodic activity identified in our study did not simply reflect event-related potentials (ERPs), we removed evoked activity from the data by subtracting ERPs before performing the time-frequency analysis. We then used linear mixed model to assess the association between the ERPs and aperiodic activity for each channel-time point. We observed R2 values ranging from 0.1 to 0.2 around stimulus onset, but otherwise correlations were negligible (Figure S14). This was observed throughout the entire time period examined. In addition, a visual comparison of the topographies of stimulus- and response-locked ERPs with those of aperiodic activity (Figure S13) revealed no clear similarities, except for a central ERP component observed around 0.25–1 second after the response, which resembled the fronto-central aperiodic component. However, while this component in aperiodic activity showed large differences between targets and non-targets, in ERPs the component differed more between 0- and 2-back conditions.

Revised version: To ensure that the aperiodic activity identified in our study did not simply reflect event-related potentials (ERPs), we removed evoked activity from the data by subtracting ERPs before performing the time-frequency analysis. We then used linear mixed model to assess the association between the ERPs and aperiodic activity for each channel-time point. We observed R2 values ranging from 0.1 to 0.2 around stimulus onset, but otherwise correlations were negligible (Figure S17). This was observed throughout the entire time period examined. In addition, a visual comparison of the topographies of stimulus- and response-locked ERPs with those of aperiodic activity (Figure S16) revealed no clear similarities, except for a central ERP component observed around 0.25–1 second after the response, which resembled the fronto-central aperiodic component. However, while this component in aperiodic activity showed large differences between targets and non-targets, in ERPs the component differed more between 0- and 2-back conditions.

Earlier version: We also repeated the analysis without subtracting the ERPs (Figure S15, Figure S16, Figure S17), and the results were largely similar. Nevertheless, when ERP was not subtracted we observed a more pronounced increase in both aperiodic and periodic power between 0 and 0.5 seconds following stimulus presentation compared to the case with subtracted ERPs.

Revised version: We also repeated the analysis without subtracting the ERPs (Figure S18, Figure S19, Figure S20), and the results were largely similar. Nevertheless, when ERP was not subtracted we observed a more pronounced increase in both aperiodic and periodic power between 0 and 0.5 seconds following stimulus presentation compared to the case with subtracted ERPs.

Replication confirms main findings

To ensure the generalisability of the findings, we repeated all analyses using two additional datasets. First we analysed a publicly available dataset from a previously published study by Nakuci et al. (2023), in which participants performed a verbal n-back task with three load levels (0-, 1-, 2-back). Next, we analysed the data from an item-recognition task.

Earlier version: In both analyses, we were able to replicate the main findings previously reported, namely the task-related modulation of periodic alpha and beta activity and the task-related changes in aperiodic activity. In the baseline-corrected data, we observed a decrease in alpha and beta power accompanied by an increase in low-frequency power (Figure S18B, Figure 10B). The latter was not evident in the grand average periodic activity (Figure S18C, Figure 10C, Figure S40) and was only detected in two participants in the n-back task (Figure S23). In the item-recognition task, peak in the theta range was present in approximately one-third of the participants (Figure S39). In half of these cases, the peaks were at 7–8 Hz, which could be considered low alpha peaks rather than theta oscillations.

Revised version: In both analyses, we were able to replicate the main findings previously reported, namely the task-related modulation of periodic alpha and beta activity and the task-related changes in aperiodic activity. In the baseline-corrected data, we observed a decrease in alpha and beta power accompanied by an increase in low-frequency power (Figure S21B, Figure 10B). The latter was not evident in the grand average periodic activity (Figure S21C, Figure 10C, Figure S43) and was only detected in two participants in the n-back task (Figure S26). In the item-recognition task, peak in the theta range was present in approximately one-third of the participants (Figure S42). In half of these cases, the peaks were at 7–8 Hz, which could be considered low alpha peaks rather than theta peaks.

Removed in the revised version: The effect of baseline correction and FOOOF decomposition on time-frequency decomposition in the item-recognition task.

Removed in the revised version: Similar to the n-back task (Figure 3, Figure S18), a decrease in alpha power is observed following instruction and stimuli presentation (up to 0.8 seconds) and continues throughout the retention period, up to 2.8 seconds post-stimulus (B). This is followed by a decrease in alpha and beta power during probe presentation, which is likely indicative of a motor response. Simultaneously, there is an increase in low-frequency power, which is most pronounced during stimulus presentation (up to 0.8 seconds) and again after probe presentation (after 2.8 seconds). The FOOOF decomposition indicates that a substantial portion of low-frequency activity could be attributed to the aperiodic component (C, D) (see also Figure S33). The data shown represent the group average over all conditions at electrode Fz, where low-frequency activity was most pronounced (see also Figure S40). Horizontal lines indicate the boundaries of the frequency ranges. See Figure S32 for the corresponding figure with a logarithmic y-axis.

Earlier version: Notably, the correlation patterns between baseline-corrected time-frequency data and FOOOF parameters, as well as between FOOOF parameters and pACF, were consistent across all three datasets. Specifically, aperiodic parameters were associated with low-frequency (theta) baseline-corrected activity, whereas periodic activity correlated with alpha and beta baseline-corrected activity and pACF (Figure S21, Figure S20, Figure S22, Figure S34, Figure S33, Figure S35).

Revised version: Notably, the correlation patterns between baseline-corrected time-frequency data and FOOOF parameters, as well as between FOOOF parameters and pACF, were consistent across all three datasets. Specifically, aperiodic parameters were associated with low-frequency (theta) baseline-corrected activity, whereas periodic activity correlated with alpha and beta baseline-corrected activity and pACF (Figure S24, Figure S23, Figure S25, Figure S37, Figure S36, Figure S38).

Earlier version: Regarding aperiodic activity, we observed frontal and parietal/occipital components in both tasks. In the itemrecognition task, these components were clearly observable only for the offset parameter, whereas for the exponent, the two components were not clearly distinguishable, likely due to lower spatial resolution (32 channels). In both tasks, aperiodic activity was modulated by task demands, with an increase in the exponent following stimuli presentation and response (Figure S27, Figure S37).

Revised version: Regarding aperiodic activity, we observed frontal and parietal/occipital components in both tasks. In the item-recognition task, these components were clearly observable only for the offset parameter, whereas for the exponent, the two components were not clearly distinguishable, likely due to lower spatial resolution (32 channels). In both tasks, aperiodic activity was modulated by task demands, with an increase in the exponent following stimuli presentation and response (Figure S30, Figure S40).

Discussion

This study provides a detailed analysis of the relationship between periodic and aperiodic neural activity during the n-back task. EEG data recorded from 57 participants performing both 0-back and 2-back tasks in visuospatial and verbal modalities were analysed to distinguish these two components within the EEG frequency spectrum. The analysis shows the presence of periodic activity in the alpha/beta frequency range prior to stimulus onset, which modulates according to task demands. Furthermore, the results suggest that low-frequency power may largely reflect the effect of baseline correction on aperiodic activity, which also exhibits its own task-related dynamics. This insight is crucial for the interpretation of EEG results, especially for continuous tasks usch as the n-back, which require sustained working memory engagement.

In addition to the interplay between the periodic and aperiodic EEG activity, the study expands on the findings of stimulus-induced aperiodic slope shift by identifying two neural components within the aperiodic activity. These aperiodic components are distinguished by their frontal and posterior topographic locations. Additionally, we observed task-dependent differences, including a slightly different slope modulation in non-target compared to target trials. However, we did not find any differences between modalities. Collectively, these findings enhance our understanding of the neural dynamics involved in cognitive processes during a continuous working memory task.

Disentangling periodic and aperiodic components in EEG activity: avoiding baseline correction

The analysis of the EEG data shows that when baseline correction is applied to the total EEG activity, changes in alpha, beta and low-frequency activity are observed following stimulus presentation compared to the prestimulus interval (Figure 3B). In particular, there is a significant increase in power in the low-frequency range after the stimulus, while a decrease in power in the alpha/beta range is observed for some time after the stimulus presentation. These findings are consistent with previous n-back studies reporting that the presentation of both target and non-target stimuli elicits long-lasting increases in low-frequency power and decreases in alpha/beta power at stimulus onset (Palomäki et al., 2012; Krause et al., 2010; Coleman et al., 2023).

Earlier version: Omitting baseline correction and decomposing the total EEG activity into periodic and aperiodic components provides a more nuanced insight into the dynamics of the spectra, showing that alpha and beta activities remain oscillatory while low-frequency (delta and theta) activity “disappears” from the periodic component. These findings suggest that the apparent power in the low-frequency range may, at least in part, reflect changes in the aperiodic component, which could contribute to the post-stimulus increase in low-frequency power observed under baseline correction in time-frequency analysis of the full EEG spectrum. Consistent with this interpretation, we found that aperiodic parameters correlated with baseline-corrected low-frequency activity. Given the current understanding of aperiodic activity, specifically based on the findings of poststimulus slope steepening in the context of increased cognitive demands (Akbarian et al., 2023b; Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022), these results suggest that changes in power within the low-frequency range may be influenced by aperiodic slope modulation rather than a distinct increase in oscillatory low-frequency activity. At the same time, alpha and beta activity remain present as part of the periodic component throughout the task.

Revised version: Omitting baseline correction and decomposing the total EEG activity into periodic and aperiodic components provides a more nuanced insight into the dynamics of the spectra, showing that alpha and beta activities are periodic while low-frequency (delta and theta) activity “disappears” from the periodic component. These findings suggest that the apparent power in the low-frequency range may, at least in part, reflect changes in the aperiodic component, which could contribute to the post-stimulus increase in low-frequency power observed under baseline correction in time-frequency analysis of the full EEG spectrum. Consistent with this interpretation, we found that aperiodic parameters exhibited the strongest correlations with low-frequency baseline-corrected activity. Given the current understanding of aperiodic activity, specifically based on the findings of post-stimulus slope steepening in the context of increased cognitive demands (Akbarian et al., 2023b; Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022), these results suggest that changes in power within the low-frequency range may be influenced by aperiodic slope modulation rather than a distinct increase in oscillatory low-frequency activity. At the same time, alpha and beta activity remain present as part of the periodic component throughout the task.

Earlier version: In line with this view, the analyses show that low-frequency activities do not necessarily exhibit an oscillatory characteristic during the n-back task (see Figure 3C, D, Figure 6). Following a detailed visual inspection at the individual level (Figure S6), we found that only two participants exhibited evident theta peak, while only one participant displayed pronounced theta and alpha peaks. The majority of participants exhibited alpha peaks, while only two participants displayed a theta peak with no substantial alpha activity. This pattern points to a low-frequency activity that may support the cognitive demands of the n-back task, yet it appears highly variable across individuals, and potentially within individuals, depending on task requirements. A recent study of visual working memory reached a similar conclusion, reporting a robust alpha peak in 90% of participants but inconsistent or absent theta peaks in 60% of participants (Van Engen et al., 2024).

Revised version: In line with this view, the analyses show that low-frequency activities do not necessarily exhibit an oscillatory characteristic during the n-back task (see Figure 3C, D, Figure 6). Following a detailed visual inspection at the individual level (Figure S9), we found that only two participants exhibited evident theta peak, while only one participant displayed pronounced theta and alpha peaks. The majority of participants exhibited alpha peaks, while only two participants displayed a theta peak with no substantial alpha activity. This pattern points to a low-frequency activity that may support the cognitive demands of the n-back task, yet it appears highly variable across individuals, and potentially within individuals, depending on task requirements. A recent study of visual working memory reached a similar conclusion, reporting a robust alpha peak in 90% of participants but inconsistent or absent theta peaks in 60% of participants (Van Engen et al., 2024).

Earlier version: A potential concern, also raised by reviewers of the original version of this paper, is that FOOOF decomposition may miss low-frequency oscillations near the edges of the spectrum. Indeed, simulations we performed showed that in scenarios with low-frequency periodic components, FOOOF may fail to identify these components, leading to inflated estimates of aperiodic parameters (see the Supplement for details). To address this issue, we performed control analyses using the recently introduced measure of rhythmicity, the phase autocorrelation function (pACF) (Myrov et al., 2024), which quantifies oscillatory activity by assessing the temporal stability of phase independent of amplitude. This metric showed no evidence of strong low-frequency periodic activity in our dataset, providing further evidence that the observed aperiodic slope reflects genuine non-oscillatory dynamics, rather than an overlooked theta peak. Furthermore, weak correlations between aperiodic parameters and pACF reinforce the independence of these aperiodic features from classical oscillations.

Revised version: A potential concern, also raised by reviewers of the original version of this paper, is that FOOOF decomposition may miss low-frequency peaks near the edges of the spectrum. Indeed, simulations we performed showed that in scenarios with low-frequency periodic components, FOOOF may fail to identify these components, leading to inflated estimates of aperiodic parameters (see the Supplement for details). To address this issue, we performed control analyses using the recently introduced measure of rhythmicity, the phase autocorrelation function (pACF) (Myrov et al., 2024), which quantifies oscillatory activity by assessing the temporal stability of phase independent of amplitude. This metric showed no evidence of strong low-frequency periodic activity in our dataset, providing further evidence that the observed aperiodic slope reflects genuine non-oscillatory dynamics, rather than an overlooked theta peak. Furthermore, weak correlations between aperiodic parameters and pACF reinforce the independence of these aperiodic features from classical oscillations.

Earlier version: The obtained results show that classical baseline correction can remove continuous oscillatory activity that is present both during baseline and after stimulus onset, because it treats all baseline signals as ‘background’ to be removed without distinguishing between transient and continuous oscillations. While this is consistent with the intended purpose of baseline correction–to highlight changes relative to ongoing activity–it may also lead to unintended consequences, such as misinterpreting aperiodic activity as an increase in post-stimulus theta oscillations. Importantly, we confirmed that these observations hold regardless of whether baseline normalisation was performed using a subtractive or divisive approach, suggesting that the conclusions are not driven by any single normalisation strategy.

Revised version: The obtained results show that classical baseline correction can remove continuous oscillatory activity that is present both during baseline and after stimulus onset, because it treats all baseline signals as ‘background’ to be removed without distinguishing between transient and continuous periodic activity. While this is consistent with the intended purpose of baseline correction– to highlight changes relative to ongoing activity–it may also lead to unintended consequences, such as misinterpreting aperiodic activity as an increase in post-stimulus theta power. Importantly, we confirmed that these observations hold regardless of whether baseline normalisation was performed using a subtractive or divisive approach, suggesting that the conclusions are not driven by any single normalisation strategy.

The modulation of periodic activity is concentrated in the alpha and beta bands

Earlier version: When decomposing the EEG signal into periodic and aperiodic components we still observed a significant attenuation in the power of alpha and beta oscillations after stimulus presentation (see Figure 3B, C and Figure 6A). This is in line with existing literature that documents the decrease of post-stimulus alpha and beta (Krause et al., 2010; Coleman et al., 2023). Moreover, the results support previously reported increased suppressive response that is elicited by target stimuli compared to non-target stimuli (Palomäki et al., 2012; Peng et al., 2015), especially in the alpha range (Figure S7).

Revised version: When decomposing the EEG signal into periodic and aperiodic components we still observed a significant attenuation in the power of alpha and beta power after stimulus presentation (see Figure 3B, C and Figure 6A). This is in line with existing literature that documents the decrease of post-stimulus alpha and beta (Krause et al., 2010; Coleman et al., 2023). Moreover, the results support previously reported increased suppressive response that is elicited by target stimuli compared to non-target stimuli (Palomäki et al., 2012; Peng et al., 2015), especially in the alpha range (Figure S10).

Our findings support existing research (Pesonen et al., 2007) by demonstrating that reductions in alpha and beta powers are more pronounced and have longer duration under increased task demands, as is typical of higher loads in n-back tasks (see Figure 7). Consistent with seminal work (Gevins et al., 1997; Pesonen et al., 2007; Krause et al., 2010), the results of this study show that working memory load mainly affects alpha/beta power in the parietal-occipital electrodes (see Figure 6B). In n-back tasks, diminished alpha power is thought to indicate increased engagement in working memory operations (Haegens et al., 2014). The functional significance of beta oscillations remains a topic of debate, with studies often linking centrally located beta oscillations to the coordination of motor responses (e.g. Pavlov and Kotchoubey, 2022; Pfurtscheller et al., 1998; Kaiser et al., 2001). However, recent research indicates that they also play a crucial role in working memory, particularly in supporting the maintenance of memorised information and preventing distractions (Schmidt et al., 2019; HajiHosseini et al., 2020). Additionally, beta oscillations are believed to facilitate the manipulation of stored information, dynamically adjusting to meet task demands and to serve a gating function by inhibiting irrelevant inputs, thus enhancing cognitive efficiency (ElShafei et al., 2022; Schmidt et al., 2019).

The strong correlation between alpha, lower beta activity and reaction times in central regions, as shown in Figure 7, suggests a shared relationship between these measures and shows that beta activity, typically associated with motor function, is most pronounced in central electrodes, although no direct temporal relationship with reaction times was observed. We also observed alpha/beta activities in the parietal electrodes, which support current theories that suggest a collaborative function for alpha and beta oscillations in supporting advanced cognitive functions such as working memory (Scharinger et al., 2017; Coleman et al., 2023; Hanslmayr et al., 2016).

Aperiodic activity dynamics and their cognitive implications

Earlier version: The analysis of aperiodic activity at each time point within the peristimulus interval builds upon prior research that has documented stimulus-induced changes in the aperiodic component, specifically the poststimulus steepening of the aperiodic slope (Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022; Kałamała et al., 2024). The results of our study indicate the presence of two components in the dynamics of the aperiodic slope during the n-back task, differentiated both temporally and spatially.

Revised version: The analysis of aperiodic activity at each time point within the peristimulus interval builds upon prior research that has documented stimulus-induced changes in the aperiodic component, specifically the post-stimulus steepening of the aperiodic slope (Gyurkovics et al., 2022; Virtue-Griffiths et al., 2022; Kałamała et al., 2024). The results of our study indicate the presence of two components in the dynamics of the aperiodic slope during the n-back task, differentiated both temporally and spatially.

Earlier version: In terms of temporality, the first component appears early in the task cycle, whereas the second component appears at a later stage (see Figure 8A). The presence of these components differs between target and non-target stimuli, with a more pronounced difference for the second component (see Figure 9). While non-target stimuli consistently exhibit a clear second component, in the case of target stimuli, the slope returns to its baseline steepness more quickly, with a significantly reduced or non-existent second slope steepening. This raises the question of whether the two components represent distinct cognitive operations inherent to the n-back task.

Revised version: In terms of temporality, the first component appears early in the task cycle, whereas the second component appears at a later stage (see Figure 8A). The presence of these components differs between target and non-target stimuli, with a more pronounced difference for the second component (see Figure 9). While non-target stimuli consistently exhibit a clear second component, in the case of target stimuli, the slope returns to its base-line steepness more quickly, with a significantly reduced or non-existent second slope steepening. This raises the question of whether the two components represent distinct cognitive operations inherent to the n-back task.

We speculate that the second component, which emerges more prominently following non-target stimuli, may be indicative of an elevated state of vigilance. Since it is rare for one target stimulus to be immediately succeeded by another, participants may become accustomed to a period of “rest” following a target stimulus, given the low anticipation for another. The increased aperiodic slope following non-target stimuli may reflect the level of alertness for potential targets, whereas after actual target stimuli, vigilance may diminish more rapidly, accounting for the absence of a second slope increase.

The two aperiodic components also exhibit distinct spatial manifestations. The first component is predominantly observed on the frontal electrodes, whereas the second component is also present in the parietal channels. This topographical distinction further supports the hypothesis that these two components reflect different cognitive processes. The activation in the parietal channels may be related to anticipatory mechanisms for the next stimulus, whereas the activation in the frontal channels may be involved in the decision-making process regarding the classification of the previous stimulus as a target.

Earlier version: Overall, the results strongly suggest that solely examining total EEG power or oscillatory dynamics, without distinctly analysing the aperiodic slope, would overlook a subtle yet potentially significant cognitive processing feature. This feature may be related to distinct working memory processes, which warrants further research.

Revised version: Overall, the results strongly suggest that solely examining total EEG power or periodic dynamics, without distinctly analysing the aperiodic slope, would overlook a subtle yet potentially significant cognitive processing feature. This feature may be related to distinct working memory processes, which warrants further research.

Earlier version: Importantly, to exclude the possibility that aperiodic activity simply reflects the event-related potentials (ERPs; see Figure S12), ERPs were subtracted before performing the time-frequency decomposition. We observed low correlations between the ERPs and aperiodic activity (Figure S14) at the stimulus onset, but no correlations after 0.5 seconds, suggesting the task demands within the n-back paradigm manifest in aperiodic modulations that are distinct from other stimulus-locked ERPs. While it is possible that response-related ERPs contributed to the second aperiodic component, the differing patterns between conditions in ERPs and aperiodic activity, along with the weak association between reaction times and aperiodic activity, suggest that ERPs and aperiodic activity capture distinct aspects of neural processing rather than reflecting the same underlying phenomenon.

Revised version: Importantly, to exclude the possibility that aperiodic activity simply reflects the event-related potentials (ERPs; see Figure S15), ERPs were subtracted before performing the time-frequency decomposition. We observed low correlations between the ERPs and aperiodic activity (Figure S17) at the stimulus onset, but no correlations after 0.5 seconds, suggesting the task demands within the n-back paradigm manifest in aperiodic modulations that are distinct from other stimulus-locked ERPs. While it is possible that response-related ERPs contributed to the second aperiodic component, several observations suggest otherwise: both aperiodic components were present throughout the entire epoch, differences between conditions diverged between ERPs and aperiodic activity (compare Figure 8 and Figure S16), and the associations with reaction times were weak.

Added in the revised version: Moreover, the distribution of the response-locked ERP qualitatively resembled the earlier frontocentral aperiodic component more than the later post-response component. Taken together, these findings suggest that ERPs and aperiodic activity capture distinct aspects of neural processing, rather than reflecting the same underlying phenomenon.

Earlier version: While contrasting response-related ERPs with aperiodic components can help address potential confounds, we believe that ERPs are not inherently separate from aperiodic or periodic activity. Instead, ERPs may reflect underlying changes in aperiodic and periodic activity. Therefore, different approaches to studying EEG activity should be seen as providing complementary rather than competing perspectives.

Revised version: While contrasting response-related ERPs with aperiodic components can help address potential confounds, we believe that ERPs are not inherently separate from aperiodic or periodic activity. For example, ERPs may reflect changes in periodic activity, such as phase resets (Makeig et al., 2002), or baseline shifts (Nikulin et al., 2007). ERPs may also capture aperiodic activity, either in the form of evoked transients triggered by an event (Shah, 2004) or induced changes in the ongoing background signal. This distinction has important implications: evoked transients can alter the broad-band spectrum without implying shifts in ongoing background activity, whereas induced aperiodic changes may signal different neural mechanisms, such as shifts in the excitation–inhibition balance (Gao et al., 2017). Therefore, different approaches to studying EEG activity should be seen as providing complementary rather than competing perspectives.

These findings are consistent with previous research that has consistently observed systematic, stimulus-induced changes in the aperiodic component that are distinct from concurrent stimulus-evoked EEG activity (Gyurkovics et al., 2022; Kałamała et al., 2024; Virtue-Griffiths et al., 2022). We thus underscore the importance of acknowledging the unique characteristics of aperiodic activity, particularly in the context of the n-back task.

Distinct functions of periodic and aperiodic EEG components during cognitive engagement

Earlier version: The findings on periodic and aperiodic EEG activities during the n-back task suggest a complex pattern of cognitive engagement. Following stimulus presentation, the analysis shows an almost immediate aperiodic slope increase in the fronto-central channels (see Figure 8), followed by a brief alpha/beta attenuation in the posterior channels (see Figure 6). Both of these features have been associated with inhibitory processes that temporarily disrupt ongoing excitatory activity to allow the brain to engage with the presented stimulus (Gyurkovics et al., 2022; Kałamała et al., 2024; Gevins et al., 1997; Pesonen et al., 2007). The results show that both the alpha/beta suppression and the aperiodic change in slope are more pronounced after target stimuli compared to non-target stimuli (see Figure 7 and Figure 9). These observations may indicate greater cognitive engagement associated with decision-making processes following the onset of the target stimulus.

Revised version: The findings on periodic and aperiodic EEG activities during the n-back task suggest a complex pattern of cognitive engagement. Following stimulus presentation, the analysis shows an almost immediate aperiodic slope increase in the fronto-central channels (see Figure 8), followed by a brief alpha/beta attenuation in the posterior channels (see Figure 6). Alpha activity has been linked to inhibitory processes; accordingly, posterior alpha (with concomitant beta) attenuation is commonly interpreted as a transient release of inhibition that facilitates sensory processing (Bonnefond and Jensen, 2012; Waldhauser et al., 2012; Wianda and Ross, 2019; Tuladhar et al., 2007). In contrast, the fronto-central increase in the aperiodic slope is compatible with a brief shift towards stronger inhibitory tone (reduced E:I) following stimulus onset (Gyurkovics et al., 2022; Kałamała et al., 2024; Gevins et al., 1997; Pesonen et al., 2007). The results show that both the alpha/beta suppression and the aperiodic change in slope are more pronounced after target stimuli compared to non-target stimuli (see Figure 7 and Figure 9). These observations may indicate greater cognitive engagement associated with decision-making processes following the onset of the target stimulus.

Further analysis of the data shows a significant difference between the 0-back and 2-back conditions in the periodic component (Figure 7), indicating more sustained alpha/beta suppression in the 2-back condition, whereas the aperiodic component did not show such spectral slope differences. Sustained suppression of periodic alpha and beta activity during the 2-back condition may reflect ongoing active engagement in the form of information encoding, retention, and updating. This finding suggests that the periodic and aperiodic components may reflect distinct cognitive processes. The characteristics of the second component in the aperiodic slope after non-target stimuli (see Figure 8), with no obvious co-occurring processes in the periodic component, further support this idea. It is important to analyse both components to ensure that the processes they reflect are not overlooked.

The results of the study show a clear distinction between the periodic and aperiodic components in relation to reaction times. These results support the association between central beta activity and motor planning (e.g. Eggermont, 2021; Pfurtscheller et al., 1998), as evidenced by its correlation with reaction times in the analysis of the periodic component (Figure 7). In contrast, the analysis indicates a relatively weak correlation between the aperiodic slope and reaction times (Figure 9). This observation differs from the most commonly reported patterns (e.g. Voytek et al., 2015; Thuwal et al., 2021; Akbarian et al., 2023b), but the evidence remains inconclusive, as some studies have reported contrasting results (e.g. Monchy et al., 2023). The discrepancy highlights the need for further research to understand how different factors influence the manifestation of the aperiodic component as a reflection of cognitive processes.

Earlier version: Prior studies have consistently reported sustained elevations in theta power during the n-back task (e.g. Palomäki et al., 2012; Krause et al., 2010; Pesonen et al., 2007). These increases, particularly in the frontal cortex, have been linked to various cognitive functions such as attention, information encoding, and cognitive control (Roux and Uhlhaas, 2014; Rossi et al., 2023; Cavanagh and Frank, 2014). While power in the theta range can indeed be oscillatory, as demonstrated in some studies, the sustained low-frequency activity observed in the analysis of scalp EEG data in this study appears to be largely influenced by aperiodic components, with post-stimulus increases potentially attributable to shifts in spectral slope rather than exclusively to oscillatory changes. Because low-frequency activity can reflect either oscillatory or aperiodic processes, distinguishing between them is key to avoiding misinterpretation.. Nevertheless, previously reported patterns of theta modulation are consistent with the notion that post-stimulus changes in aperiodic activity may reflect enhanced inhibitory control commensurate with task demands (Voytek et al., 2015; Gao et al., 2017).

Revised version: Prior studies have consistently reported sustained elevations in theta power during the n-back task (e.g. Palomäki et al., 2012; Krause et al., 2010; Pesonen et al., 2007). These increases, particularly in the frontal cortex, have been linked to various cognitive functions such as attention, information encoding, and cognitive control (Roux and Uhlhaas, 2014; Rossi et al., 2023; Cavanagh and Frank, 2014). While power in the theta range can indeed be oscillatory, as demonstrated in some studies, the sustained low-frequency activity observed in the analysis of scalp EEG data in this study appears to be largely influenced by aperiodic components, with post-stimulus increases potentially attributable to shifts in spectral slope rather than exclusively changes in periodic activity. Because low-frequency activity can reflect either periodic or aperiodic processes, distinguishing between them is key to avoiding misinterpretation. Nevertheless, previously reported patterns of theta modulation are consistent with the notion that post-stimulus changes in aperiodic activity may reflect enhanced inhibitory control commensurate with task demands (Voytek et al., 2015; Gao et al., 2017).

Main findings replicated in additional datasets

Earlier version: Analyses of control data from a previously published study employing n-back task (Nakuci et al., 2023) confirmed our initial results, indicating that similar task-related modulations in alpha and beta periodic activities, as well as aperiodic activity dynamics, can be divided into frontal and parietal components. However, although both aperiodic components were observed in the control dataset, their temporal differentiation was less pronounced (Figure S27, Figure S29). Notably, in this comparative dataset, the aperiodic slope was only related to task load, with a steeper slope observed for the 2-back level, and no significant differences based on stimulus type (Figure S28). Furthermore, in the comparative dataset, there was an association with reaction times in the parietal channels, more so than in our original dataset.

Revised version: Analyses of control data from a previously published study employing n-back task (Nakuci et al., 2023) confirmed our initial results, indicating that similar task-related modulations in alpha and beta periodic activities, as well as aperiodic activity dynamics, can be divided into frontal and parietal components. However, although both aperiodic components were observed in the control dataset, their temporal differentiation was less pronounced (Figure S30, Figure S32). Notably, in this comparative dataset, the aperiodic slope was only related to task load, with a steeper slope observed for the 2-back level, and no significant differences based on stimulus type (Figure S31). Furthermore, in the comparative dataset, there was an association with reaction times in the parietal channels, more so than in our original dataset.

To test whether these findings generalise across different task paradigms, we applied the same methods to an item-recognition task, which showed similar low-frequency dynamics despite having distinct mechanisms of memory maintenance and a non-continuous structure (Figure 10). Analyses of the original n-back dataset, the control n-back dataset, and the item-recognition dataset all point to the same core pattern: low-frequency power in baseline-corrected data correlates with aperiodic activity, while periodic activity is associated with alpha/beta power in both baseline-corrected data and pACF. This cross-task consistency underscores the robustness of the interpretation that periodic and aperiodic components should be carefully distinguished when studying working memory.

Theoretical and methodological implications

Earlier version: This study highlights the importance of refined EEG analysis methods that take into account the dynamic nature of both periodic and aperiodic components during cognitive tasks. The findings suggest a need for reconsidering baseline correction techniques, particularly in continuous paradigms where traditional baseline subtraction can obscure meaningful neural dynamics by potentially misrepresenting aperiodic activity as low-frequency (theta) oscillations (Gyurkovics et al., 2021). Such misrepresentation is particularly relevant under sustained cognitive demand, as corrections intended to normalise data may inadvertently lead to misleading or artefactual interpretations of neural activity. Further investigation is warranted to explore how changes in the aperiodic slope might underlie the commonly observed increases in frontal low-frequency power, which have traditionally been documented as theta oscillations. In order to further validate this interpretation, future research should extend to a wider range of working memory and cognitive control tasks.

Revised version: This study highlights the importance of refined EEG analysis methods that take into account the dynamic nature of both periodic and aperiodic components during cognitive tasks. The findings suggest a need for reconsidering baseline correction techniques, particularly in continuous paradigms where traditional baseline subtraction can obscure meaningful neural dynamics by potentially misrepresenting aperiodic activity as low-frequency (theta) power (Gyurkovics et al., 2021). Such misrepresentation is particularly relevant under sustained cognitive demand, as corrections intended to normalise data may inadvertently lead to misleading or artefactual interpretations of neural activity. Further investigation is warranted to explore how changes in the aperiodic slope might underlie the commonly observed increases in frontal low-frequency power, which have traditionally been referred to as theta oscillations. In order to further validate this interpretation, future research should extend to a wider range of working memory and cognitive control tasks.

Added in the revised version: The aperiodic slope dynamics observed during the Moreover, the distribution of the response-locked ERP qualitatively resembled the earlier frontocentral aperiodic component more than the later post-response component. Taken together, these findings suggest that ERPs and aperiodic activity capture distinct aspects of neural processing, rather than reflecting the same underlying phenomenon.

Added in the revised version: While contrasting response-related ERPs with aperiodic components can help address potential confounds, we believe that ERPs are not inherently separate from aperiodic or periodic activity. For example, ERPs may reflect changes in periodic activity, such as phase resets (Makeig et al., 2002), or baseline shifts (Nikulin et al., 2007). ERPs may also capture aperiodic activity, either in the form of evoked transients triggered by an event (Shah, 2004) or induced changes in the ongoing background signal. This distinction has important implications: evoked transients can alter the broad-band spectrum without implying shifts in ongoing background activity, whereas induced aperiodic changes may signal different neural mechanisms, such as shifts in the excitation–inhibition balance (Gao et al., 2017). Therefore, different approaches to studying EEG activity should be seen as providing complementary rather than competing perspectives.

Added in the revised version: These findings are consistent with previous research that has consistently observed systematic, stimulus-induced changes in the aperiodic component that are distinct from concurrent stimulus-evoked EEG activity (Gyurkovics et al., 2022; Kałamała et al., 2024; Virtue-Griffiths et al., 2022). We thus underscore the importance of acknowledging the unique characteristics of aperiodic activity, particularly in the context of the n-back task.

Added in the revised version: Distinct functions of periodic and aperiodic EEG components during cognitive engagement

Added in the revised version: The findings on periodic and aperiodic EEG activities during the n-back task suggest a complex pattern of cognitive engagement. Following stimulus presentation, the analysis shows an almost immediate aperiodic slope increase in the fronto-central channels (see Figure 8), followed by a brief alpha/beta attenuation in the posterior channels (see Figure 6). Alpha activity has been linked to inhibitory processes; accordingly, posterior alpha (with concomitant beta) attenuation is commonly interpreted as a transient release of inhibition that facilitates sensory processing (Bonnefond and Jensen, 2012; Waldhauser et al., 2012; Wianda and Ross, 2019; Tuladhar et al., 2007). In contrast, the fronto-central increase in the aperiodic slope is compatible with a brief shift towards stronger inhibitory tone (reduced E:I) following stimulus onset (Gyurkovics et al., 2022; Kałamała et al., 2024; Gevins et al., 1997; Pesonen et al., 2007). The results show that both the alpha/beta suppression and the aperiodic change in slope are more pronounced after target stimuli compared to non-target stimuli (see Figure 7 and Figure 9). These observations n-back task in this study align with the E:I balance framework (Gao et al., 2017; Vogels and Abbott, 2009; Lim and Goldman, 2013). Specifically, stimulus-induced shifts in the slope likely reflect a proportional modulation of ongoing excitatory activities to meet processing demands (Gyurkovics et al., 2022; Voytek et al., 2015; Virtue-Griffiths et al., 2022), consistent with a transient increase in inhibitory control that relates to frontal theta indices of cognitive control (e.g. Scharinger et al., 2017; Cavanagh and Frank, 2014; Ratcliffe et al., 2022). Building on this account, lifespan studies show that the aperiodic exponent becomes flatter with advancing age at rest. This is often described as greater neural noise, and this shift relates to lower cognitive performance and prospective decline (Voytek et al., 2015; Cesnaite et al., 2023; Finley et al., 2024). In this context, age is a plausible moderator of task evoked slope dynamics in the n-back task, with older adults expected to show reduced or delayed steepening. This is consistent with event-related evidence regarding aperiodic modulation and resting-state links to cognition (Kałamała et al., 2024; Finley et al., 2024; Cesnaite et al., 2023; Dave et al., 2018). As our main sample did not consist of young adults, it is possible that age moderates these aperiodic slope effects. However, as the present study was not designed or powered to test for interactions between age and condition, we treat this as a theoretical implication to be evaluated in future work involving a broader age range. Taken together, these observations suggest that aperiodic slope modulations should be considered a potential biomarker of cognitive control, and that their relationship with established measures should be investigated to predict performance across tasks.

Removed in the revised version: The aperiodic slope dynamics observed during the n-back task in this study align with the E:I balance framework (Gao et al., 2017; Vogels and Abbott, 2009; Lim and Goldman, 2013). Specifically, the results suggest that stimulus-induced shifts in the slope likely reflect a proportional modulation of ongoing excitatory activities tailored to meet processing demands (Gyurkovics et al., 2022; Voytek et al., 2015; Virtue-Griffiths et al., 2022). These adjustments in the post-stimulus aperiodic component may signify enhanced inhibitory control, which has implications for cognitive control capacities previously associated with changes in the frontal theta activity (e.g. Scharinger et al., 2017; Cavanagh and Frank, 2014; Ratcliffe et al., 2022). The findings provide support for the hypothesis that modifications in aperiodic activity following cognitive load could serve as a reliable index of cognitive control mechanisms. Future studies should consider aperiodic slope modulations as a potential biomarker for cognitive control, investigating how these measures correlate with traditional indicators and whether they can predict cognitive performance across different task contexts.

Removed in the revised version: To better understand the functional role of the observed second component in the aperiodic activity (second slope steepening) following non-target trials, a study design incorporating a larger dataset with an increased number of target-followed-by-target sequences is essential. This would enable a detailed comparison of aperiodic components in scenarios where a target is succeeded by another target versus when it is followed by a non-target, potentially revealing how attentional mechanisms are reset or sustained under different conditions.

Earlier version: Furthermore, investigating the potential relationship between aperiodic slope dynamics and theta oscillations in tasks requiring enhanced cognitive control could provide deeper insights into the neural basis of cognitive flexibility and decision-making. It is also imperative to consider individual variations in theta dynamics and their relationship with aperiodic slope modulation, as this could highlight potential biomarkers for cognitive efficiency or susceptibility to cognitive fatigue.

Revised version: Furthermore, investigating the potential relationship between aperiodic slope dynamics and theta oscillations in tasks that require enhanced cognitive control could provide a deeper understanding of the neural basis of cognitive flexibility and decision-making. Considering individual variations in theta dynamics and their relationship with aperiodic slope modulation is also important, as this could reveal potential biomarkers for cognitive efficiency or susceptibility to cognitive fatigue.

Added in the revised version: To better understand the functional role of the second aperiodic component observed after non-target trials (second slope steepening), future work should use a design with a larger dataset and a greater number of sequences in which a target is followed by another target. This would enable a detailed comparison of aperiodic components in scenarios where a target is succeeded by another target versus when it is followed by a non-target, potentially revealing how attentional mechanisms are reset or sustained under different conditions.

Earlier version: Finally, the findings highlight a potential issue for EEG-based cross-frequency coupling analyses, particularly phase-amplitude coupling (Tort et al., 2010; Van Der Meij et al., 2012; Tseng et al., 2019). Numerous studies have demonstrated gamma amplitude modulation by the phase of theta (e.g. Park et al., 2013; Köster et al., 2014; Goodman et al., 2018), posited to organise information in working memory (e.g. Rajji et al., 2016; Brooks et al., 2020; McGill and Kieffaber, 2024; Good- man et al., 2018). However, the results of this study suggest that scalp-recorded low-frequency fluctuations may be predominantly influenced by aperiodic shifts rather than true oscillations. Consequently, it may be unwarranted to assume an inherently oscillatory theta band when investigating theta-gamma coupling. In future cross-frequency coupling studies, it is therefore crucial to confirm that theta power reflects genuine oscillations, rather than aperiodic activity, before drawing conclusions about cross-frequency interactions.

Revised version: Finally, the findings highlight a potential issue for EEG-based cross-frequency coupling analyses, particularly phase-amplitude coupling (Tort et al., 2010; Van Der Meij et al., 2012; Tseng et al., 2019). Numerous studies have demonstrated gamma amplitude modulation by the phase of theta (e.g. Park et al., 2013; Köster et al., 2014; Goodman et al., 2018), posited to organise information in working memory (e.g. Rajji et al., 2016; Brooks et al., 2020; McGill and Kieffaber, 2024; Goodman et al., 2018). However, the results of this study suggest that scalp-recorded low-frequency fluctuations may be predominantly influenced by aperiodic shifts rather than true oscillations. Consequently, it may be unwarranted to assume an inherently oscillatory theta band when investigating theta-gamma coupling. In future cross-frequency coupling studies, it is therefore crucial to confirm that theta power reflects genuine oscillations, rather than aperiodic activity, before drawing conclusions about cross-frequency interactions.

Strengths and limitations

The main strength of the study is the replication of key findings by performing the same analyses on two independent samples. Replication of the results in an independent sample of younger individuals performing a standard verbal n-back task extends the generalisability of the findings beyond the original cohort of older adults. Furthermore, replication of these results in the item-recognition task demonstrates the robustness of conclusions across different task paradigms.

Furthermore, we used pACF, an alternative measure of oscillations, to verify the robustness of the approach, thereby ensuring that low-frequency oscillatory activity was not simply missed by FOOOF near the edges of the spectrum.

There are a number of limitations that need to be considered. First, the task difficulty in the 2-back condition may have been too high. Some participants had to be excluded from the analysis due to poor performance. This exclusion may have reduced the range of cognitive abilities in the sample of older adults, potentially inflating them.

Earlier version: In addition, the higher level of difficulty in the tasks may have influenced the observed differences between the 0-back and 2-back conditions, potentially affecting the magnitude and direction of the reported neural modulations. For the sample of older adults, a simplified 2-back task using a single modality may have been more appropriate. The single modality (verbal, with letters displayed in the centre of the screen) may help to reduce task complexity and align better with the cognitive abilities of older participants. In addition, the dual-modality task imposed time constraints that limited the number of trials for each condition, potentially leading to participant fatigue and reduced data quality.

Revised version: In addition, the higher level of difficulty in the tasks may have influenced the observed differences between the 0-back and 2-back conditions, potentially affecting the magnitude and direction of the reported neural modulations. For the sample of older adults, a simplified 2-back task using a single modality may have been more appropriate. The single modality (verbal, with letters displayed in the centre of the screen) may help to reduce task complexity and align better with the cognitive abilities of older participants. In addition, the dualmodality task imposed time constraints that limited the number of trials for each condition, potentially leading to participant fatigue and reduced data quality.

Conclusions

Earlier version: The study highlights the critical importance of distinguishing between the periodic and aperiodic components of the EEG spectrum in order to ensure accurate interpretation of changes traditionally attributed to oscillatory activity. Using the n-back task, we have demonstrated that the aperiodic component, which is often dismissed as mere background noise, reflects important features of cognitive processing. The findings also challenge the routine practice of baseline correction, which can obscure crucial modulations of aperiodic activity, particularly in continuous cognitive tasks. It is important to note that a substantial proportion of the low-frequency activity commonly observed in scalp EEG may in fact be attributable to shifts in the aperiodic slope. It is therefore crucial to independently verify whether the observed low-frequency signals are genuinely oscillatory or primarily aperiodic. This is of particular importance for accurately deciphering the neural underpinnings of working memory and cognitive control, and highlights the need for a refined approach in the analysis and interpretation of EEG data.

Revised version: The study highlights the critical importance of distinguishing between the periodic and aperiodic components of the EEG spectrum in order to ensure accurate interpretation of changes traditionally attributed to oscillatory activity. Using the n-back task, we have demonstrated that the aperiodic component, which is often dismissed as mere background noise, reflects important features of cognitive processing. The findings also challenge the routine practice of baseline correction, which can obscure crucial modulations of aperiodic activity, particularly in continuous cognitive tasks. It is important to note that a substantial proportion of the low-frequency activity commonly observed in scalp EEG may in fact be attributable to shifts in the aperiodic slope. It is therefore crucial to independently verify whether the observed low-frequency signals are genuinely periodic (oscillatory) or primarily aperiodic. This is of particular importance for accurately deciphering the neural underpinnings of working memory and cognitive control, and highlights the need for a refined approach in the analysis and interpretation of EEG data.

Methods

Participants

We analysed EEG recordings from fifty-seven volunteers (nine males) aged 55–78 years (M = 67.4, SD = 6.4), who initially participated in a study on subjective cognitive complaints. All participants had normal or corrected-to-normal vision and were free from neurological or psychiatric illness or cognitive decline, as assessed by neuropsychological testing prior to study entry. Of these participants, twenty-seven reported subjective cognitive complaints, while thirty did not. The current study focused exclusively on within-subject comparisons and did not investigate subjective cognitive complaints or between-group differences. We excluded data from 10 participants from all analyses due to poor quality.

The study protocol was approved by the Medical Ethics Committee of the Republic of Slovenia (protocol number 0120-128/2019/9; 21.08.2019). All study participants gave written informed consent prior to the initiation of any study-related procedures.

Experimental task and procedure

The participants were instructed to press the left button when the letter or position matched the letter/position n-back trials before (i.e. on target trials), and the right button otherwise (on non-target trials). Some of the stimuli were lures, i.e. stimuli that would have been considered targets in the 1-back condition but appeared during the 2-back condition.

Earlier version: The participants completed the n-back task in two modalities (visuospatial and verbal) with two load conditions (0- and 2-back) (Figure 1). During the task, participants were presented with white capital letters on a grey background one at a time in one of nine possible positions, in a continuous sequence. Each letter was presented on the screen for 250 ms, followed by a 3000 ms interval during which participants were required to respond. In the verbal task conditions, participants were instructed to memorise the presented letter, whereas in the visuospatial task conditions, participants were asked to memorise the position of the presented letter. The order of modality was counterbalanced across participants.

Revised version: The participants completed the n-back task in two modalities (visuospatial and verbal) with two load conditions (0- and 2-back) (Figure 1). During the task, participants were presented with white capital letters on a grey background one at a time in one of nine possible positions, in a continuous sequence. Each letter was presented on the screen for 250 ms, followed by a 3000 ms interval during which participants were required to respond. In the verbal task conditions, participants were instructed to memorise the presented letter, whereas in the visuospatial task conditions, participants were asked to memorise the position of the presented letter.

Added in the revised version: Schematic representation of the n-back task.

Added in the revised version: The n-back working memory task was conducted in two distinct modalities: visuospatial and verbal. In the 2-back condition, participants had to identify whether the current stimulus matched the one presented two steps previously. In the visuospatial modality, the target was a spatial location, whereas in the verbal modality, the target was a letter. In the 0-back condition participants’ task was to respond to a predefined target, with the type of target corresponding to the task modality. Target stimuli are highlighted in orange.

Added in the revised version: Schematic of the analysis.

Added in the revised version: The time domain data were first transformed into the time-frequency domain by convolution with superlets (Moca et al., 2021). Next, the periodic and aperiodic components of the power spectrum density were estimated for each time point using FOOOF (Fitting Oscillations and One-Over-F) algorithm (Donoghue et al., 2020b). The aperiodic component was characterised by the aperiodic slope (the negative counterpart of the exponent parameter) and the offset, which together describe the underlying broadband spectral shape.

Added in the revised version: Effect of baseline correction and FOOOF decomposition on time-frequency decomposition in the n-back task.

Added in the revised version: Illustration of the effect of baseline correction and FOOOF decomposition on EEG time-frequency analysis. (A) “Raw” total power from time-frequency decomposition is difficult to interpret due to the 1/f power scaling of EEG power spectra (here representing average data across subjects, channels and conditions). In panel B we applied baseline correction (decibel conversion) using a pre-stimulus interval of -0.5 to -0.2 seconds for comparison. The baseline correction showed a significant decrease in alpha and beta power from 0 to 1 second and a concomitant increase in low-frequency power lasting up to 2 seconds. The observed changes were very similar across different choices of baseline correction (Figure S2). However, it’s unclear from the baseline-corrected data whether the observed changes in the low-frequency range reflect periodic or aperiodic contributions. To disentangle these components, we decomposed the time-frequency signal into periodic and aperiodic contributions using spectral parameterisation. (C) The periodic component includes only the parameterised spectral peaks, reconstructed from Gaussian fits to the power spectrum (see Methods for details). (D) The aperiodic component reflects the 1/f-like background activity. This decomposition suggests that changes in power in the low-frequency (delta and theta) range may largely reflect changes in aperiodic activity. See Figure S1 for the corresponding figure with a logarithmic y-axis.

Added in the revised version: Correlations between baseline-corrected time-frequency and FOOOF-decomposed EEG activity

Added in the revised version: For each FOOOF parameter, we estimated its similarity to baseline-corrected time-frequency EEG activity at each channel-frequency-timepoint using a linear mixed model. We observed strong correlations between periodic activity and baseline-corrected time-frequency activity in the alpha and beta ranges. In contrast, the exponent and offset showed strong correlations in the low-frequency range, with the exponent also showing negative correlations in the gamma range. These results suggest that changes in aperiodic activity appear as low-frequency power in baseline-corrected time-frequency plots. Note that while R2 is strictly non-negative, we assigned its sign based on the beta coefficient from the fixed effect in the model to facilitate interpretation. The values presented are averaged across all channels; see Figure S3 for topographical distributions. Condition-specific correlations are shown in Figure S4, Figure S5, and Figure S6.

Added in the revised version: Phase autocorrelation function and its correlations with FOOOF parameters.

Added in the revised version: Phase autocorrelation function (pACF) (Myrov et al., 2024) is an alternative measure of oscillations that is sensitive to rhythmicity rather than the amplitude. Patterns of pACF (A) were largely consistent with periodic activity (B), particularly in the alpha range, supporting the notion that pACF and periodic activity reflect the same underlying processes. In contrast, the correlations with exponent and offset were low (only results for exponent are shown as results for exponent and offset were essentially the same). See Figure S7 and Figure S8 for detailed figures.

Added in the revised version: Changes in periodic activity as a function of time.

Added in the revised version: (A) The inspection of periodic activity revealed strong activity in the alpha and beta frequency bands, with a sharp decrease at around 0.5 seconds post-stimulus. The power was stronger in the 0-back condition, compared to the 2-back condition (see also Figure 7). (B) Early beta activity was most prominent at occipital and frontal channels. Note that colour scale ranges in panel B differ between frequency ranges.

Added in the revised version: Results of the linear mixed model analysis of periodic activity for comparison between conditions.

Added in the revised version: Significant values are highlighted on the heatmap and marked with yellow circles on the topographies. Only factors with significant differences are shown. A p-value was interpreted as significant if it was significant in at least 3 channels or 3 time points (see Methods for details). Note that there are small discrepancies between significant values on heatmaps and topographies due to averaging across time or channels. A significant difference was observed between the 2-back and 0-back conditions, with a reduction in activity in the 2-back condition, particularly in the alpha and beta frequency bands, starting at 0.5 seconds post-stimulus. The differences between stimulus types were most evident from 0.3 to 1 second post-stimulus, with decreased activity for targets compared to non-targets across the whole scalp. Additionally, a smaller effect of the modality × stimulus type interaction was observed from 0.6 to 1 second post-stimulus (see also Figure S10 for detailed visualisation of the interaction). Associations with reaction times were significant in the beta band across central channels. These associations exhibited a positive correlation in the early phase (0 to 0.5 seconds) and a negative correlation in the later phase (1 to 1.5 seconds).

Added in the revised version: Changes in aperiodic activity (exponent, interpreted as aperiodic slope) as a function of time.

Added in the revised version: (A) We averaged the time course of the aperiodic activity (exponent) over all channels and observed two components (peaks) of the aperiodic activity. The first component peaked around 0.3 s post-stimulus in the frontal channels. The second component peaked at 0.7 s, was stronger in parietal channels and differed between non-target and target conditions (see Figure 9). The time course of the offset parameter was comparable, although the separation between the frontal and parietal components was more pronounced (see Figure S11). The vertical lines in panel A represent mean reaction time modes. The shaded areas represent 95% Cousineau-Morey within-subjects confidence intervals, adjusted so that non-overlapping intervals correspond to statistically significant differences (Cousineau, 2017).

Added in the revised version: Results of the linear mixed model analysis of aperiodic exponent (interpreted as slope) for comparison between conditions.

Added in the revised version: Significant values are shown in blue on line plots and marked with yellow circles on topographies. Only factors with significant differences are shown. The only significant differences between conditions were observed between target and non-target conditions, where the exponent was higher for targets in the early phase (0 to 0.5 s post-stimulus) and for non-targets in the middle phase (0.5 to 1 s post-stimulus). There was also a small association with reaction times in occipital channels around 1.5 seconds post-stimulus. Results were similar for the offset parameter, with an additional effect of load and n-back × stimulus type interaction (Figure S14).

Added in the revised version: The effect of baseline correction and FOOOF decomposition on time-frequency decomposition in the item-recognition task.

Added in the revised version: Similar to the n-back task (Figure 3, Figure S21), a decrease in alpha power is observed following instruction and stimuli presentation (up to 0.8 seconds) and continues throughout the retention period, up to 2.8 seconds post-stimulus (B). This is followed by a decrease in alpha and beta power during probe presentation, which is likely indicative of a motor response. Simultaneously, there is an increase in low-frequency power, which is most pronounced during stimulus presentation (up to 0.8 seconds) and again after probe presentation (after 2.8 seconds). The FOOOF decomposition indicates that a substantial portion of low-frequency activity could be attributed to the aperiodic component (C, D) (see also Figure S36). The data shown represent the group average over all conditions at electrode Fz, where low-frequency activity was most pronounced (see also Figure S43). Horizontal lines indicate the boundaries of the frequency ranges. See Figure S35 for the corresponding figure with a logarithmic y-axis.

Added in the revised version: The order of modality was counterbalanced across participants.

Each modality always began with the 0-back load condition in which the target stimulus was presented for 10 seconds at the beginning of the block and stayed the same throughout the whole block. The 0-back load condition was followed by a 2-back load condition of the same modality, where participants had to remember each stimulus as it appeared and discern whether it matched the one displayed two trials prior. The 0-back and 2-back load conditions each comprised three 5-minute blocks of continuously presented stimuli. For each modality, the 0-back load condition comprised a total of 201 stimuli, 39 of which were targets, and the 2-back load condition comprised a total of 240 stimuli, 39 of which were targets, and 30 were lures. All stimuli were presented in a pseudorandomised order.

The task began only after participants indicated their readiness by pressing a designated key, after reviewing the instructions and clarifying any doubts with the experimenter. Before each condition, participants completed a practice session consisting of a short block to familiarise themselves with the task demands. The session included a variation of the main task in which participants practised identifying both the target letter and its location, as well as responding to the 2-back challenge for both letters and positions. During the main task, response accuracy and reaction times were recorded by keystroke on the Cedrus response pad (Cedrus Corporation, San Pedro, CA, USA, model RB-540). Participants were instructed to favour accurate responses over fast responses. The participants’ responses were monitored, and only trials with correct responses were included in the EEG analysis.

The stimuli were presented in Helvetica font on a 24-inch LCD screen with a refresh rate of 120 Hz. The stimuli, arranged in a 3 × 3 matrix, spanned approximately 7.22 degrees of both vertical and horizontal visual angles. The task was programmed and executed using PsychoPy 3 (Peirce et al., 2022).

EEG data acquisition

Earlier version: We recorded EEG activity using either a BrainAmp MR-plus or ActiChamp amplifier (Brain Products GmbH, Germany). We acquired the recordings with a 64-channel ActiCap plus a reference channel, configured according to the extended 10-10 international system layout. To accommodate all participants, we used two slightly different electrode layouts, differing only in the position of two electrodes. Due to inconsistencies in electrode placement between subjects, we excluded channels AFz, Iz, PO9, and PO10 from the analysis, resulting in a dataset of 63 electrodes (including the reference channel). We grounded the electrode setup at AFz, with FCz serving as the reference. During the EEG recording, we band-pass filtered the signal between 0.1 and 250 (or 1000) Hz with a slope of 12dB/octave and digitised it at a sampling rate of 500, 1000, or 2500 Hz. After time-frequency decomposition, we downsampled the data to 100 Hz.

Revised version: We recorded EEG activity using either a BrainAmp MR-plus or ActiChamp amplifier (Brain Products GmbH, Germany). We acquired the recordings with a 64-channel ActiCap plus a reference channel, configured according to the extended 10-10 international system layout. To accommodate all participants, we used two slightly different electrode layouts, differing only in the position of two electrodes. Due to inconsistencies in electrode placement between subjects, we excluded channels AFz, Iz, PO9, and PO10 from the analysis, resulting in a dataset of 63 electrodes (including the reference channel). We grounded the electrode setup at AFz, with FCz serving as the reference. During the EEG recording, we band-pass filtered the signal between 0.1 and 250 (or 1000) Hz with a slope of 12dB/octave and digitised it at a sampling rate of 500, 1000, or 2500 Hz.

Added in the revised version: After time-frequency decomposition, we downsampled the data to 100 Hz.

EEG data preprocessing

Earlier version: We preprocessed the EEG data using custom MAT-LAB scripts (MathWorks, USA) that integrated functions from EEGLAB (Delorme and Makeig, 2004) and ER-PLAB (Lopez-Calderon and Luck, 2014). High-pass filtering at 0.1 Hz was performed using a finite impulse response (FIR) filter (slope: 12 dB/octave). A CleanLine filter was then used to reduce power line noise at 50 Hz (Mullen, 2012). Next, we segmented the data into epochs ranging from -1.5 to 2.8 seconds, time-locked to stimulus presentations, followed by baseline correction to remove residual signal drift. Epochs and channels were manually inspected to identify and exclude segments with excessive artefacts. Independent component analysis decomposition was performed using the AMICA algorithm (Palmer et al., 2011), applied to the data bandpass filtered between 1 and 60 Hz; the solution was later applied to the non-bandpass filtered data. We manually inspected the independent components to identify and remove those representing blinks, horizontal eye movements, muscle noise, line noise, and other artefacts. After the ICA, channel time series reconstruction and channel interpolation were performed, and the data were referenced to the average reference.

Revised version: We preprocessed the EEG data using custom MAT-LAB scripts (MathWorks, USA) that integrated functions from EEGLAB (Delorme and Makeig, 2004) and ER-PLAB (Lopez-Calderon and Luck, 2014). High-pass filtering at 0.1 Hz was performed using a finite impulse response (FIR) filter (slope: 12 dB/octave). A Clean-Line filter was then used to reduce power line noise at 50 Hz (Mullen, 2012). Next, we segmented the data into epochs ranging from -1.5 to 2.8 seconds, time-locked to stimulus presentations, followed by baseline correction to remove residual signal drift. Epochs and channels were manually inspected to identify and exclude segments with excessive artefacts. Independent component analysis decomposition was performed using the AMICA algorithm (Palmer et al., 2011), applied to the data bandpass filtered between 1 and 60 Hz; the solution was later applied to the non-bandpass filtered data. We manually inspected the independent components to identify and remove those representing blinks, horizontal eye movements, muscle noise, line noise, and other artefacts. After the ICA, channel time series reconstruction and channel interpolation were performed, and the data were referenced to the average reference.

Subsequent EEG analyses were conducted only on subjects and conditions that met the required standards, defined as having a minimum of 30 good epochs after cleaning and an accuracy rate of at least 0.6 within a given condition. Exclusion criteria were applied at the level of individual conditions within subjects, rather than excluding entire participants from the analyses. Consequently, only 21% and 10% of participants remained for the visuospatial and verbal lure conditions, respectively. As a result, the lure condition was excluded from all subsequent analyses.

Estimation of periodic and aperiodic component

First, we excluded trials with incorrect or missing responses from all subsequent analyses. ERPs related to early sensory potentials can affect estimates of aperiodic activity (Gyurkovics et al., 2022), therefore, average ERPs were removed prior to time-frequency decomposition, as has been previously done in similar studies (Gyurkovics et al., 2022; Kałamała et al., 2024). This involved computing ERPs for each subject/condition and then subtracting these ERPs from single-trial data for each subject/condition. For completeness, we also analysed the data without subtracted ERPs.

To estimate both periodic and aperiodic components of EEG activity, we performed time-frequency decomposition using superlets, with a range of frequencies from 3 to 50 Hz, in steps of 1 Hz (width of the base wavelet: 3 cycles, order: linearly spaced from 1 to 20, multiplicative combination of wavelets) as implemented in Fieldtrip (Oostenveld et al., 2011). Superlets, a variant of the Morlet wavelet transformation, involve multiple decompositions with varying cycles, the results of which are combined using a geometric mean to achieve enhanced temporal and spectral resolution, known as ‘super-resolution’ (Moca et al., 2021). The results of the time-frequency decomposition were averaged across trials for each condition and participant, and then decimated to 100 Hz in order to save space and facilitate further analysis.

Earlier version: We then estimated the periodic and aperiodic components from the decomposed data for each time point, condition, and participant using the FOOOF algorithm from the Python specparam package (Donoghue et al., 2020b). To avoid edge artefacts, FOOOF parameters were estimated for a period between -0.5 and 2 seconds around the stimulus. The FOOOF algorithm parameterises the frequency spectrum in a semi-logarithmic space, where only power values are log-transformed, using a set of parameters separately for periodic and aperiodic activity. Specifically, neural power spectra (NPS) over a set of frequencies (F) are described as the sum of the aperiodic and periodic components:

Revised version: We then estimated the periodic and aperiodic components from the decomposed data for each time point, condition, and participant using the FOOOF algorithm from the Python specparam package (version 2.0.0rc2) (Donoghue et al., 2020b). To avoid edge artefacts, FOOOF parameters were estimated for a period between -0.5 and 2 seconds around the stimulus. The FOOOF algorithm parameterises the frequency spectrum in a semi-logarithmic space, where only power values are log-transformed, using a set of parameters separately for periodic and aperiodic activity. Specifically, neural power spectra (NPS) over a set of frequencies (F) are described as the sum of the aperiodic and periodic components:

The aperiodic component is modelled as where b is the offset, k is the knee (bend), and χ is the aperiodic exponent (sometimes referred to as the spectral slope).

The periodic component is modelled as a sum of Gaussians:

Each Gaussian is described by three parameters: c, a, and w, representing center frequency, power, and bandwidth, respectively.

Earlier version: Fitting was performed using default specparam parameters: peak width limits = 0.5–12.0 Hz, maximum number of peaks to fit = Inf, minimum peak height = 0, peak threshold = 2.0, aperiodic mode = ‘fixed’ (i.e. the knee parameter was set to 0). Note that the FOOOF fitting was performed on trial-averaged data to increase the stability of the parameter estimates. In contrast, Gyurkovics et al. (2022), who also investigated taskrelated changes in aperiodic activity, stabilised their fits by estimating power spectra over 1-second time windows at the single-trial level. Although effective, this strategy involved a trade-off in temporal resolution that our trial-averaging approach seeks to avoid.

Revised version: Fitting was performed using default specparam parameters: peak width limits = 0.5–12.0 Hz, maximum number of peaks to fit = Inf, minimum peak height = 0, peak threshold = 2.0, aperiodic mode = ‘fixed’ (i.e. the knee parameter was set to 0). Note that the FOOOF fitting was performed on trial-averaged data to increase the stability of the parameter estimates. In contrast, Gyurkovics et al. (2022), who also investigated task-related changes in aperiodic activity, stabilised their fits by estimating power spectra over 1-second time windows at the single-trial level. Although effective, this strategy involved a trade-off in temporal resolution that our trial-averaging approach seeks to avoid.

Earlier version: To assess the fit of the FOOOF models, we examined R2 and model error metrics (Figure S9). Model error was defined as the absolute difference between predicted and actual values. However, as we have shown in the simulation (see section Control analyses below), model fit metrics can be misleading. By simulating data and estimating parameters, we observed cases where the model fit was high, but the estimated parameters deviated substantially from the ground truth values. This highlights the limitations of relying solely on fit metrics to assess parameter accuracy.

Revised version: To assess the fit of the FOOOF models, we examined R2 and model error metrics (Figure S12). Model error was defined as the absolute difference between predicted and actual values. However, as we have shown in the simulation (see section Control analyses below), model fit metrics can be misleading. By simulating data and estimating parameters, we observed cases where the model fit was high, but the estimated parameters deviated substantially from the ground truth values. This highlights the limitations of relying solely on fit metrics to assess parameter accuracy.

Earlier version: In addition, we examined the number of peaks identified per model. On average, 1.9–2 peaks were identified, showing high consistency and indicating that the models did not overfit by detecting an unrealistic number of peaks (Figure S10). We also repeated the analysis with the maximum number of peaks set to three, which gave virtually identical results (results not shown).

Revised version: In addition, we examined the number of peaks identified per model. On average, 1.9–2 peaks were identified, showing high consistency and indicating that the models did not overfit by detecting an unrealistic number of peaks (Figure S13). We also repeated the analysis with the maximum number of peaks set to three, which gave virtually identical results (results not shown). The outcomes of the analyses conducted with the off-set and exponent parameters were found to be highly comparable. Consequently, the main text only includes analyses using the exponent parameter, whereas the Supplement contains analyses using the offset parameter.

Removed in the revised version: The outcomes of the analyses conducted with the offset and exponent parameters were found to be highly comparable. Consequently, the main text only includes analyses using the exponent parameter, whereas the Supplement contains analyses using the offset parameter.

Earlier version: To allow comparisons with full time-frequency power and to facilitate comparisons across participants and conditions, we reconstructed the periodic component in a semi-log space based on the estimated parameters for each participant and condition using the formula in Equation 3. This reconstruction includes only the parameterised oscillatory peaks and is referred to as ‘periodic power’ or ‘periodic activity’, a measure we used in all subsequent analyses. Importantly, the reconstructed periodic components were strictly non-negative: although Gaussians asymptotically approach zero, they never reach it, resulting in smooth transitions without explicit zeros between peaks.

Revised version: To allow comparisons with full time-frequency power and to facilitate comparisons across participants and conditions, we reconstructed the periodic component in a semi-log space based on the estimated parameters for each participant and condition using the formula in Equation 3. This reconstruction includes only the parameterised spectral peaks and is referred to as ‘periodic power’ or ‘periodic activity’, a measure we used in all subsequent analyses. Importantly, the reconstructed periodic components were strictly non-negative: although Gaussians asymptotically approach zero, they never reach it, resulting in smooth transitions without explicit zeros between peaks.

We visually assessed the time-frequency plots, both in their original form and after baseline correction, before performing formal statistical analyses. To assess the effect of baseline correction on the results, we compared several baseline correction methods, including: (a) decibel conversion (10 ∗ log10(data/baseline)), (b) relative change ((data − baseline)/baseline), (c) normalised change ((data − baseline)/(data + baseline)), (d) absolute change (data − baseline). In addition, we considered three baseline periods (relative to stimulus onset): (a) from -500 to -200 ms, (b) from -300 to 0 ms, (b) from -500 to 0 ms.

Comparison of baseline-corrected data and FOOOF parameters

We compared baseline-corrected time-frequency activity with both periodic and aperiodic components. To better understand their relationship, we first visually inspected grand-average plots and then performed mass univariate comparisons. Specifically, we examined the relationship between baseline-corrected power and FOOOF parameters for each channel, time, and frequency point.

A simple approach to assessing this relationship would be to compute correlations across subjects at each point. However, this approach would violate the independence assumption of simple linear models, as each subject contributed multiple observations across different experimental conditions. A more appropriate approach was to use a linear mixed model that included both subject and condition as random effects.

Although condition is typically treated as a fixed effect, in this case, we were not interested in the specific conditions themselves. Instead, we treated them as a random sample from a larger population of possible conditions. Baseline-corrected activity was treated as the dependent variable, while periodic or aperiodic activity served as a fixed effect predictor variable.

Earlier version: For each test, we estimated the marginal R2 (Nak- agawa and Schielzeth, 2013), a measure of variance explained for linear mixed models, analogous to R2 in general linear models. The marginal R2 is defined as variance explained by the fixed factors using the equation: where represents the variance explained by the fixed effects, and correspond to variance of random effects (subject- and condition-related variance, respectively), and represents the residual variance. The fixed-effects variance can be computed as the variance of the predicted values.

Revised version: For each test, we estimated the marginal R2 (Nakagawa and Schielzeth, 2013), a measure of variance explained for linear mixed models, analogous to R2 in general linear models. The marginal R2 is defined as variance explained by the fixed factors using the equation: where represents the variance explained by the fixed effects, and correspond to variance of random effects (subject- and condition-related variance, respectively), and represents the residual variance. The fixed-effects variance can be computed as the variance of the predicted values.

Although R2 is strictly non-negative, we assigned the sign based on the beta coefficient from the fixed effect in the mixed model.

Earlier version: Additionally, to assess variability across participants, we visualised time-averaged periodic activity for each participant (Figure S6).

Revised version: Additionally, to assess variability across participants, we visualised time-averaged periodic activity for each participant (Figure S9).

Validation of FOOOF results

As reviewers of the original manuscript raised concerns that FOOOF might miss peaks near the edges of the spectrum (Gerster et al., 2022), we performed two control analyses to test the validity of the FOOOF results.

Earlier version: First, we assessed whether the FOOOF decomposition can accurately resolve low-frequency periodic activity by simulating ground truth data with systematically varied parameter values. Details of this simulation are provided in the Supplement. Briefly, we found that low-frequency periodic components were often conflated with aperiodic components, resulting in undetected periodic activity at very low frequencies. This issue was more pronounced when fitting was restricted to frequencies ≥ 3 Hz (as in our analyses of experimental data). In addition, although model fit metrics were generally good, they did not always indicate accurate parameter estimation. These results suggest that even when true low-frequency periodic activity is present, FOOOF may fail to detect it. Conversely, changes in the aperiodic parameters estimates may actually be a consequence of low-frequency periodic activity.

Revised version: First, we assessed whether the FOOOF decomposition can accurately resolve low-frequency periodic activity by simulating ground truth data with systematically varied parameter values. Details of this simulation are provided in the Supplement. Briefly, we found that low-frequency periodic components were often conflated with aperiodic components, resulting in undetected periodic activity at very low frequencies. This issue was more pronounced when fitting was restricted to frequencies ≥ 3 Hz (as in our analyses of experimental data). In addition, although model fit metrics were generally good, they did not always indicate accurate parameter estimation. These results suggest that even when true low-frequency periodic activity is present, FOOOF may fail to detect it. Conversely, changes in the aperiodic parameters estimates may actually be a consequence of low-frequency periodic activity. Importantly, as noted by the reviewer, this behaviour reflects an explicit design choice of the algorithm: to avoid overfitting ambiguous peaks at the edges of the spectrum, FOOOF excludes peaks that are too close to the boundaries. This exclusion is controlled by the _bw_std_edge parameter, which defines the distance that a peak must be from the edge in order to be retained (in units of standard deviation; set to 1.0 by default). Therefore, although the algorithm is functioning as intended, users should exercise be careful interpreting aperiodic parameters in datasets where low-frequency periodic activity might be expected.

Earlier version: Second, we used an alternative method to assess whether signals at different frequencies were oscillatory. Specifically, we used the recently developed phase autocorrelation function (pACF) (Myrov et al., 2024). Unlike amplitude-based methods such as FOOOF, pACF is sensitive to the ‘rhytmicity’ or temporal stability of the phase of a signal. pACF achieves this by calculating the cross-spectrum between a complex signal and its time-delayed copy, then normalising to unit magnitude to remove amplitude information. Because pACF isolates phase relationships while being inherently independent of amplitude, it serves as a complementary measure to FOOOF. Consistent results between FOOOF and pACF would strengthen our conclusion that low-frequency (theta) oscillations were truly absent from our recordings, rather than simply undetectable by FOOOF.

Revised version: Second, we used an alternative method to assess whether signals at different frequencies were oscillatory. Specifically, we used the recently developed phase autocorrelation function (pACF) (Myrov et al., 2024). Unlike amplitude-based methods such as FOOOF, pACF is sensitive to the ‘rhythmicity’ or temporal stability of the phase of a signal. pACF achieves this by calculating the cross-spectrum between a complex signal and its time-delayed copy, then normalising to unit magnitude to remove amplitude information. Because pACF isolates phase relationships while being inherently independent of amplitude, it serves as a complementary measure to FOOOF. Consistent results between FOOOF and pACF would strengthen our conclusion that low-frequency (theta) oscillations were truly absent from our recordings, rather than simply undetectable by FOOOF.

Earlier version: Sample pACF is estimated using formula: where CSx,x(l) represents a normalised crossspectrum between a complex signal X and its delayed copy at lag l, and N is the total number of samples.

Revised version: Sample pACF is estimated using formula: where CSx,x(l) represents a normalised cross-spectrum between a complex signal X and its delayed copy at lag l, and N is the total number of samples.

To allow comparison with the time-resolved FOOOF decomposition, the time-resolved pACF was estimated as follows: where t is a timestep, l is a lag, W is a window size (in samples), N is the number of lags, and z is the phase difference between a signal and its delayed version. The size of the moving window was set to 2.5 cycles, with lags ranging from 1 to 3 cycles in 0.1 increments. These parameters were chosen based on Myrov et al. (2024). The moving window size was chosen to ensure stability of the results without excessive smoothing. The selected lag range was based on the consideration that too small lags are predominantly influenced by filtering effects, whereas large lags may capture activity from non-local sources (V. Myrov, personal communication, January 21, 2025).

To compute pACF, the signals were convolved with complex Morlet wavelets (width: 3 cycles). We did not use superlets here, as they were primarily designed for power estimation rather than phase analysis. In fact, when we tested superlets, the pACF values increased linearly with frequency – probably due to the way superlets combine wavelets of different widths, which can affect the phase autocorrelation of the convolved signal. To ensure an accurate comparison between FOOOF and pACF, we repeated the FOOOF decomposition using the same Morlet wavelets, which gave results comparable to those obtained with superlets, particularly at lower frequencies.

To directly compare FOOOF and pACF results, we used linear mixed models, following the same approach used for the comparison between baseline-corrected activity and FOOOF. In this analysis, pACF served as the dependent variable, while FOOOF parameters were included as fixed effect predictors, with condition modelled as a random effect. We fitted a separate model for each FOOOF parameter and estimated the marginal R2.

Comparisons between experimental conditions

We performed statistical comparisons between conditions using linear mixed models with the fitlme function in MATLAB. An advantage of using linear mixed models is their ability to accommodate missing data, which improves parameter estimates through partial pooling. We fitted the models as follows:

Here, y refers to the exponent or offset in the case of aperiodic activity and to the log-power in the case of periodic activity. We included reaction times (rt) in the model to control for differences between conditions and to assess the relationship between motor response and EEG activity.

Earlier version: To estimate representative reaction times for each participant and condition, we removed reaction times above a z-value of 3 and those below 200 ms within each participant and condition. We then fitted an exGaussian distribution to estimate the parameter µ using the maximum likelihood method as implemented in the retimes package in R (Massidda, 2013). The exGaussian distribution, a positively skewed distribution, is useful for modelling reaction time distributions (Balota and Yap, 2011), and the parameter µ can be interpreted as the mode of the distribution. Other parameters (e.g. σ, τ) were excluded to avoid overfitting and because our aim was to obtain a simple, robust estimate of the central tendency rather than to model the entire distribution.

Revised version: To estimate representative reaction times for each participant and condition, we removed reaction times above a z-value of 3 and those below 200 ms within each participant and condition. We then fitted an exGaussian distribution to estimate the parameter µ using the maximum likelihood method as implemented in the retimes package in R (Massidda, 2013). The ex-Gaussian distribution, a positively skewed distribution, is useful for modelling reaction time distributions (Balota and Yap, 2011), and the parameter µ can be interpreted as the mode of the distribution. Other parameters (e.g. σ, τ) were excluded to avoid overfitting and because our aim was to obtain a simple, robust estimate of the central tendency rather than to model the entire distribution.

We fitted models in a mass univariate manner, that is for each channel, frequency (where applicable), and time point separately.

Fixed effect were estimated using the maximum likelihood method, and the degrees of freedom were approximated using the Satterthwaite equation. We corrected the p-values using the Benjamini-Yekutieli false discovery rate (FDR) correction (Benjamini and Yekutieli, 2001). Given that under the false discovery rate (FDR) correction, a q proportion of significant p-values are false positives, and that EEG effects occur in clusters (in time, frequency, and/or space) (Van Ede and Maris, 2016), we additionally considered as significant only clusters of p-values that spanned either at least three time points or at least three channels. We set the significance threshold at α = .05.

For the purposes of visualisation, p-values were averaged across channels (for heatmaps or lines) or across time (for topographies). Consequently, there are some discrepancies between heatmaps/lines and topographies.

Comparison of aperiodic activity with event-related potentials

Earlier version: To rule out the possibility that aperiodic activity merely reflected evoked potentials, we first removed evoked activity before time-frequency decomposition (as described above). We then estimated associations between event-related potentials (ERPs) and aperiodic activity. Specifically, we used linear mixed models to assess this association at each channel-time point, with ERP activity as a predictor of aperiodic activity and condition included as a random effect, similar to the comparison between FOOOF parameters and baselinecorrected data.

Revised version: To rule out the possibility that aperiodic activity merely reflected evoked potentials, we first removed evoked activity before time-frequency decomposition (as described above). We then estimated associations between event-related potentials (ERPs) and aperiodic activity. Specifically, we used linear mixed models to assess this association at each channel-time point, with ERP activity as a predictor of aperiodic activity and condition included as a random effect, similar to the comparison between FOOOF parameters and baseline-corrected data.

In addition, we visually compared the topographies of aperiodic activity with those of stimulus-locked and response-locked ERPs.

The results are shown in the Supplement.

Replication: n-back task

To test the generalisability of the results, we performed two replication analyses using two different EEG datasets: an n-back task and an item-recognition task. For the n-back task, we used an EEG dataset from a previously published study by Nakuci et al. (2023). This dataset includes data from 21 subjects (17 female, mean age 42 ± 12 years) who performed a verbal n-back task (0-, 1-, and 2-back) while being recorded with a 256-channel HydroCel Geodesic Sensor Net. The task involved 150 trials per condition, 50 of which were targets. Each stimulus was presented for 400 ms, followed by an interstimulus interval of 2000 ms.

The dataset was used in its preprocessed form, as the unprocessed data was not available. The preprocessing included band-pass filtering, artifact subspace reconstruction (ASR) (Mullen et al., 2013), independent component analysis, channel interpolation, another round of ASR, re-referencing to the average reference, epoching from -1000 to 2000 ms around stimulus presentation, and baseline correction using the prestimulus period (Nakuci et al., 2023). Notably, the data were band-pass filtered between 0.5 and 50 Hz, so we restricted the time-frequency decomposition and FOOOF analysis to frequencies between 3 and 35 Hz.

We carried out all other analytical steps, including time-frequency decomposition, FOOOF analysis, and statistical comparisons, in the same manner as in the primary analysis.

The results of these analyses are presented in the Supplement.

Replication: item-recognition task

To test generalisability across tasks, we analysed EEG recordings from another study, which included 48 participants (10 male, mean age 23.8 ± 5 years). All participants were right-handed, had normal or corrected-to-normal vision, and were free of neurological or psychiatric illness. Most participants were psychology students at the University of Ljubljana, while the rest responded to invitations sent via email and social media. All participants provided informed consent, approved by the Ethics Committee of the Faculty of Arts, University of Ljubljana.

In this study, we used a Sternberg item-recognition task. In this paradigm, a series of visual target stimuli is presented, followed by a probe stimulus that may or may not be identical to one of the previous targets. Participants’ task is to indicate whether the probe (test) stimulus was part of the initial set (Sternberg, 1969).

Earlier version: The experiment included three task conditions performed under two working memory loads: (1) presentation of two or four stimuli that differed in colour but were displayed in the same location; (2) presentation of two or four stimuli that differed in position but were of the same colour (black); (3) presentation of two or four stimuli that differed in both colour and position (integration). Each condition consisted of 432 trials, with 216 trials per working memory load, divided into 18 blocks of 48 trials each. Participants were randomly assigned to three groups, each performing the task in two of the three conditions, resulting in 32 participants per condition and 16 participants per comparison of conditions. A schematic diagram and a detailed description of the task can be found in Figure S31.

Revised version: The experiment included three task conditions performed under two working memory loads: (1) presentation of two or four stimuli that differed in colour but were displayed in the same location; (2) presentation of two or four stimuli that differed in position but were of the same colour (black); (3) presentation of two or four stimuli that differed in both colour and position (integration). Each condition consisted of 432 trials, with 216 trials per working memory load, divided into 18 blocks of 48 trials each. Participants were randomly assigned to three groups, each performing the task in two of the three conditions, resulting in 32 participants per condition and 16 participants per comparison of conditions. A schematic diagram and a detailed description of the task can be found in Figure S34.

Stimuli were presented on a 27-inch LCD screen with a refresh rate of 120 Hz. The stimuli were presented on a grey background in the centre of the screen, within a frame 250 pixels wide and 250 pixels high. The stimuli appeared as circles 20 pixels in diameter. When different positions were required, the stimuli were spaced at a distance equal to six times the stimulus radius. The frame spanned approximately 4.05 degrees of vertical and horizontal visual angle, and the diameter of each stimulus spanned approximately 0.32 degrees of visual angle. The task was programmed and ran using PsychoPy 2 (Peirce et al., 2022).

For the EEG measurements we used a BrainAmp EEG amplifier (Brain Products GmbH, Germany) and a 32-channel ActiCAP EEG cap, configured according to the international 10/20 system. Participants provided responses using a Cedrus response pad, RB series, model 530. We monitored horizontal eye movements using a camera and a dedicated HEOG channel ([F7— F8]/2). We manually removed trials in which participants tracked the target stimuli with their eyes.

We preprocessed the EEG data using the same pipeline as for the n-back task, except that the data were segmented into epochs ranging from -1700 to 5400 miliseconds, time-locked to the cue presentations. Subsequent EEG analyses were performed in a manner identical to the n-back task, including only epochs with correct responses.